# Week 7 Problem 1

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/info490-fa16/blob/master/Week2/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
%matplotlib inline

# plotting tools
import matplotlib as mpl
import matplotlib.pyplot as plt
import numpy as np
import seaborn as sns

# testing tools
from nose.tools import (
    assert_equal, assert_is_instance, assert_is_not,
    assert_greater_equal, assert_less_equal
)
from numpy.testing import assert_array_almost_equal
```

# Problem 1.

In this problem, you will be creating a function that plots the sine and cosine functions (sin(x) and cos(x)) on the same plot.  Here is a sample plot:
![sine cosine function plot](sine_cos.png "Sine and Cosine Plot")

The main purpose of this problem is for you to become familiarized with 
making simple plots in matplotlib.  Before you start, you should take note 
of a few important things:

- You don't have to make your plot look exactly like this example. If your plot looks visually OK, and if the test code cell doesn't produce any errors, your solution is correct.
- Note that code for creating a Figure object and an Axes object is already provided: `fig, ax = plt.subplots()`
- Also note that the plot_sin_cos function returns an instance (named ax) of the Axes object. You should use ax to create your plot (check out the lesson notebook for lesson 2 for examples on how to do this). The reason we write our function this way to return an Axes instance is because it can be used to extract data values for testing.
- Your plot should have a title and axis lables.
- In order for your function to be graded correctly, you need to first plot the sine function, and then the cosine function.  If you do not do this, you will not pass the last few assertion tests.
- The $x$-axis should be in radians (not degrees) and should at least go from $-\pi$ to $+\pi$.
- Since seaborn has been imported, you are allowed to use it in this problem.


<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
def plot_sin_cos():
    """
    Plots the sine function.
    
    Parameters
    ----------
    None
    
    Returns
    -------
    A matplotlib.Axes instance.
    """
    
    fig, ax = plt.subplots()

    # YOUR CODE HERE
    
    return ax
```

```python
ax = plot_sin_cos()
```

<!-- grade_id: cube_test, points: 30, grade: true, solution: false -->
```python
assert_is_instance(
    ax, mpl.axes.Axes,
    msg="Your function should return a matplotlib.axes.Axes object."
)

assert_equal(len(ax.lines), 2)

assert_is_not(
    len(ax.title.get_text()), 0,
    msg="Your plot doesn't have a title."
)
assert_is_not(
    ax.xaxis.get_label_text(), '',
    msg="Change the x-axis label to something more descriptive."
)
assert_is_not(
    ax.yaxis.get_label_text(), '',
    msg="Change the y-axis label to something more descriptive."
)
    
x, y1 = ax.lines[0].get_xydata().T
assert_less_equal(x.min(), -3.14)
assert_greater_equal(x.max(), 3.14)
assert_array_almost_equal(y1, np.sin(x))

x, y2 = ax.lines[1].get_xydata().T
assert_less_equal(x.min(), -3.14)
assert_greater_equal(x.max(), 3.14)
assert_array_almost_equal(y2, np.cos(x))
```

If you are not passing the last two sets of assertion tests, did you make sure that you are plotting the sine function first and then the cosine function?  