# Week 3 Problem 3

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/info490-fa16/blob/master/Week2/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
from nose.tools import assert_equal
```

# Problem 1. Decision Tree

In this problem, you will use functions and conditionals to build a [decision tree](https://en.wikipedia.org/wiki/Decision_tree) that helps you decide whether to go to class or stay at home.
Decision trees are used in [decision tree learning](https://en.wikipedia.org/wiki/Decision_tree_learning),
which we will learn more about in the later part of this course when we introduce [machine learning](https://en.wikipedia.org/wiki/Machine_learning).

```python
from IPython.display import Image
Image(filename='decision_tree.png', width=300, height=300) 
```

We will write a function for each question (i.e., each node). We will then combine the two questions into a function named `i_want_class()` that represents this simple decision tree, which return `True` if the input values predicts that you would like to go to class and `False` if you prefer to stay at home.

## 1. Function: `is_weekday()`

In the following code cell, write a function named `is_weekday()` that represents the top node, "_is today a weekday?_" The function takes one argument (an integer named `day` between 0 and 6), and returns `True` if `day` is between 1 and 5 and `False` otherwise. You do not need to handle inputs other than integers from 0 to 6.

<!-- grade_id: p311_answer, points: 0, grade: false, solution: true -->
```python
def is_weekday(day):
    # YOUR CODE HERE
```

<!-- grade_id: p311_test, points: 5, grade: true, solution: false -->
```python
assert_equal(is_weekday(0), False)
assert_equal(is_weekday(1), True)
assert_equal(is_weekday(2), True)
assert_equal(is_weekday(3), True)
assert_equal(is_weekday(4), True)
assert_equal(is_weekday(5), True)
assert_equal(is_weekday(6), False)
```

## 2. Function: `is_math()`

In the following code cell, write a function named `is_math()` that represents the second node, "_is today's subject math?_" The function takes one argument (a string named `subject` that is either `"math"` or `"literature"`), and returns `True` if `subject` is `"math"` and `False` if `subject` is `"literature"`. You do not need to handle inputs other than `"math"` and `"literature"`.

<!-- grade_id: p312_answer, points: 0, grade: false, solution: true -->
```python
def is_math(subject):
    # YOUR CODE HERE
```

<!-- grade_id: p312_test, points: 5, grade: true, solution: false -->
```python
assert_equal(is_math('math'), True)
assert_equal(is_math('literature'), False)
```

## 3. Function: `i_want_class()`

In the following code cell, use the two functions, `is_weekday()` and `is_math()`, inside a function named `i_want_class()` to represent the decision tree, "_do I want to go to class today?_" The function takes an integer named `day` as the first argument, and a string named `subject` as the second argument. It returns `True` if you decide to go to class and `False` if you decide to stay at home.

<!-- grade_id: p313_answer, points: 0, grade: false, solution: true -->
```python
def i_want_class(day, subject):
    # YOUR CODE HERE
```

<!-- grade_id: p313_test, points: 5, grade: true, solution: false -->
```python
assert_equal(i_want_class(0, 'math'), False)
assert_equal(i_want_class(0, 'literature'), False)
assert_equal(i_want_class(1, 'math'), True)
assert_equal(i_want_class(1, 'literature'), False)
assert_equal(i_want_class(2, 'math'), True)
assert_equal(i_want_class(3, 'literature'), False)
assert_equal(i_want_class(5, 'math'), True)
assert_equal(i_want_class(5, 'literature'), False)
```

# Problem 2. 

This problem helps you practice conditionals, iterations, and list comprehensions.

## 1. Function: get_informatics()

Use conditionals and the [str()](https://docs.python.org/3/library/stdtypes.html#str) function to write a function named `get_informatics()` that takes a positive integer between 1 and 100 (inclusive) as the input argument, and converts the integer to string type. The convertion follows the rules below:

- For any input integer, use `str()` to convert it to string type, and return the converted string. For example, if the input is 1 (`<class 'int'>`) then the function should return "1" (`<class 'str'>`).
- If the integer is a multiple of 3, return "info" instead.
- If the integer is a multiple of 7, return "matics" instead.
- If the integer is a multiple of both 3 and 7, return "informatics" instead.

Please be careful about the spelling and the case of the string that `get_informatics()` returns. Note that you do not need to deal with inputs other than integers between 1 and 100.

<!-- grade_id: p321_answer, points: 0, grade: false, solution: true -->
```python
def get_informatics(number):
    '''
    Convert the input integer to string type.
    
    Parameters
    __________
    input: An integer between 1 and 100, inclusive.
    
    Returns
    _______
    output: A string.
    '''
    
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: p321_test, points: 8, grade: true, solution: false -->
```python
assert_equal(get_informatics(1), '1')
assert_equal(get_informatics(2), '2')
assert_equal(get_informatics(3), 'info')
assert_equal(get_informatics(4), '4')
assert_equal(get_informatics(5), '5')
assert_equal(get_informatics(6), 'info')
assert_equal(get_informatics(7), 'matics')
assert_equal(get_informatics(8), '8')
assert_equal(get_informatics(9), 'info')
assert_equal(get_informatics(10), '10')
assert_equal(get_informatics(11), '11')
assert_equal(get_informatics(12), 'info')
assert_equal(get_informatics(13), '13')
assert_equal(get_informatics(14), 'matics')
assert_equal(get_informatics(15), 'info')
assert_equal(get_informatics(16), '16')
assert_equal(get_informatics(17), '17')
assert_equal(get_informatics(18), 'info')
assert_equal(get_informatics(19), '19')
assert_equal(get_informatics(20), '20')
assert_equal(get_informatics(21), 'informatics')
assert_equal(get_informatics(22), '22')
assert_equal(get_informatics(23), '23')
assert_equal(get_informatics(24), 'info')
assert_equal(get_informatics(25), '25')
assert_equal(get_informatics(26), '26')
assert_equal(get_informatics(27), 'info')
assert_equal(get_informatics(28), 'matics')
assert_equal(get_informatics(29), '29')
assert_equal(get_informatics(30), 'info')
assert_equal(get_informatics(31), '31')
assert_equal(get_informatics(32), '32')
assert_equal(get_informatics(33), 'info')
assert_equal(get_informatics(34), '34')
assert_equal(get_informatics(35), 'matics')
assert_equal(get_informatics(36), 'info')
assert_equal(get_informatics(37), '37')
assert_equal(get_informatics(38), '38')
assert_equal(get_informatics(39), 'info')
assert_equal(get_informatics(40), '40')
assert_equal(get_informatics(41), '41')
assert_equal(get_informatics(42), 'informatics')
assert_equal(get_informatics(43), '43')
assert_equal(get_informatics(44), '44')
assert_equal(get_informatics(45), 'info')
assert_equal(get_informatics(46), '46')
assert_equal(get_informatics(47), '47')
assert_equal(get_informatics(48), 'info')
assert_equal(get_informatics(49), 'matics')
assert_equal(get_informatics(50), '50')
```

## 2. List Comprehension

In the following code cell, use the `get_informatics()` function in a list comprehension to create a list named `data`. In the list comprehension, iterate through integers between 1 and 100 (inclusive), convert every integer to string type according to the rules in `get_informatics()`, and add the resulting strings to the `data` list. You may be familiar with the `list.append(element)` function, but note that you cannot use that function inside list comprehensions. 

<!-- grade_id: p322_answer, points: 0, grade: false, solution: true -->
```python
data = [
    # YOUR CODE HERE
]
```

<!-- grade_id: p322_test, points: 7, grade: true, solution: false -->
```python
assert_equal(len(data), 100)
assert_equal(data, ['1', '2', 'info', '4', '5', 'info', 'matics', '8', 'info', '10', 
                    '11', 'info', '13', 'matics', 'info', '16', '17', 'info', '19', '20', 
                    'informatics', '22', '23', 'info', '25', '26', 'info', 'matics', '29', 'info', 
                    '31', '32', 'info', '34', 'matics', 'info', '37', '38', 'info', '40', 
                    '41', 'informatics', '43', '44', 'info', '46', '47', 'info', 'matics', '50', 
                    'info', '52', '53', 'info', '55', 'matics', 'info', '58', '59', 'info', 
                    '61', '62', 'informatics', '64', '65', 'info', '67', '68', 'info', 'matics', 
                    '71', 'info', '73', '74', 'info', '76', 'matics', 'info', '79', '80', 'info', 
                    '82', '83', 'informatics', '85', '86', 'info', '88', '89', 'info', 'matics', 
                    '92', 'info', '94', '95', 'info', '97', 'matics', 'info', '100'])
```

```python

```