# Week 3 Problem 2

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/accy570-fa16/blob/master/Week1/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
from nose.tools import assert_equal
from nose.tools import assert_almost_equal
import math
```

# Problem 1.

The main purpose of this problem is to familiarize yourself with creating function definitions.  For this first problem, I want you to create a function called `in_radians` that takes the argument `degrees`, which will be an angle expressed in degrees.  In this function, you will need to convert the input `degrees` into radians.  Once this conversion is done, you will need to return the angle in radians that you calculated.

You should notice that the math module is already imported, so you should consider using that in your definition below.

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
def in_radians(degrees):
    # YOUR CODE HERE
```

<!-- grade_id: degree_test, points: 5, grade: true, solution: false -->
```python
assert_almost_equal(in_radians(180), 3.14159265358)
assert_almost_equal(in_radians(36), 0.6283185307)
assert_almost_equal(in_radians(215), 3.75245789178)
assert_almost_equal(in_radians(0), 0.0)
```

# Problem 2.

Now that we have created a function that converts an angle in units of degrees into an angle in units of radians, we will now put it to good use. In the following function definition, you will create a function called `sine_of` that takes the argument `degrees`.  Inside of this argument you should call the function `in_radians`, and your function should return the sine of the angle.

Once again, notice that the math module is imported; a function inside this module might prove very helpful in this problem!

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
def sine_of(degrees):
    # YOUR CODE HERE
```

<!-- grade_id: sine_test, points: 5, grade: true, solution: false -->
```python
assert_almost_equal(sine_of(45),0.7071067811865475)
assert_almost_equal(sine_of(180),0.0)
assert_almost_equal(sine_of(60),0.8660254037844386)
```

# Problem 3.

Finally, imagine that you have a right triangle.  You know the length of the hypotenuse as well as one angle (in degrees), but you want to know the length of the side opposite the angle that you know.  In order to find this length of the opposite triangle leg, you must find the sine of the angle you know and then multiply it by the hypotenuse (as you should know from trigonometry).  

In the following function called `find_opp`, it will take the arguments `degrees` and `hypo` (short for hypotenuse length). Inside of the function, it will multiply the sine of the angle `degrees` by the hypotenuse and then return the resulting number.

Note: this will be a very short and simple function definition if you call the definitions you made above!

<!-- grade_id: p3_answer, points: 0, grade: false, solution: true -->
```python
def find_opp(degrees, hypo):
    # YOUR CODE HERE
```

<!-- grade_id: SOH_test, points: 10, grade: true, solution: false -->
```python
assert_almost_equal(find_opp(45,2),1.414213562373095)
assert_almost_equal(find_opp(30,4),1.9999999999999998)
assert_almost_equal(find_opp(60,1),0.8660254037844386)
```