# Week 3 Problem 1

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

***

This problem will focus on Chapters 9 and 10 of [TLCL](https://sourceforge.net/projects/linuxcommand/?source=dlp) such as **file permissions** and **process monitoring**.

```python
from nose.tools import assert_equal
```

Let's start by making a test directory called `w3p1_testdir`. If you want to start over, simply rerun the following code chunk to recreate the directory. As a general tip, if you find yourself running errant code a few times in a row and/or things get in the wrong place, simply rerun the following cell to start over. For this assignment to work properly, all the commands have to be executed properly **in succession**, so there is a lot of room for error. Please try to understand what each command you issue is doing so you can keep track of what has happened within `w3p1_testdir`. You may find it helpful to open a terminal window to explore the filesystem while you work through this notebook.

```python
%%bash

# set a name for the test directory
dir='w3p1_testdir'

# make a directory for testing
# remove it and remake if it already exists
if [[ -e ~/$dir ]]; then
    rm -rf ~/$dir 
fi

mkdir ~/$dir
```

We can check the permissions/mode of the file like so:

```python
ls_res = !ls -l ~ | grep w3p1_testdir 
print(ls_res[0][:10])
```

The first line grabs the whole `ls -l` entry for the home directory and the second line displays the first 10 characters of the entry, called the *file attributes*.

# Problem 1.

The purpose of this problem is to give you experience creating directories with certain permissions. This will be an extremely useful skill if you find yourself working on a shared Linux server and need to share files with a colleague. In the following cell, use the `%%bash` magic to exectute a few linux commands in succession that 

1. Make a directory inside `~/w3p1_testdir` called `subdir`
2. Change the permissions of that directory to `rwxrw----` (octal 760)

**Note**: The `mkdir` command will flag an error if you try to create a directory that already exists. This may happen to you if you run the script you'll write below more than once. You can either use the `-p` flag with `mkdir` which says "make a directory here unless it already exists", remove the directory manually and try again, or rerun the initial script above to start everything fresh.

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
%%bash 
 # YOUR CODE HERE
```

The output of the following cell should look like: 

```
total 4
drwxrw---- 2 data_scientist users 4096 Aug 31 19:03 subdir
```

```python
!ls -l /home/data_scientist/w3p1_testdir/
```

<!-- grade_id: p1_test, points: 10, grade: true, solution: false -->
```python
p1_res = !ls -l /home/data_scientist/w3p1_testdir/ | grep subdir
assert_equal(p1_res[0][:10], 'drwxrw----')
```

# Problem 2.

The purpose of this problem is to give you experience checking and correcting permissions issues. The following cell will get a [script](https://github.com/UI-DataScience/info490-fa16/blob/master/Week3/assignments/scripts/w3p1.sh) from the course github repo then list the contents of the directory we created. The script itself is supposed to create a lot of subdirectories within `w3p1_testdir/subdir`, one for each pokemon in [this text file](https://github.com/cervoise/pentest-scripts/blob/master/password-cracking/wordlists/pokemon-list-en.txt). It will try to change directories into `w3p1_testdir/subdir`, so that had better exist upon running it. 

```python
%%bash
cd ~/w3p1_testdir/subdir/
wget https://raw.githubusercontent.com/UI-DataScience/info490-fa16/master/Week3/assignments/scripts/w3p1.sh
ls -l 
```

To run the script we downloaded, all we have to do is write out the file path like so:

```python
! /home/data_scientist/w3p1_testdir/subdir/w3p1.sh
```

But there is a problem. Your task is to figure out what the problem is and solve it to get the script to run. Write your code below to solve the problem and run the script. The autograder will check how many files are in the `w3p1_testdir/subdir` directory. There should be 721 (719 Pokemon, `w3p1.sh`, and `pokemon-list-en.txt`). If your directory has some other number of files or is not composed by the files mentioned above, you will likely not receive full points.

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
%%bash 
 # YOUR CODE HERE
```

<!-- grade_id: p2_test, points: 10, grade: true, solution: false -->
```python
file_count = !ls /home/data_scientist/w3p1_testdir/subdir/ | wc -l
assert_equal(file_count, ['721'])
```

As a reminder, this notebook must run start to finish without error. While you were working, it is possible you performed some set of commands that will not be reproduced when our autograder reruns your work. To ensure you get the points for this assignment, we highly recommend you run the notebook once more from start to finish by selecting `Kernel -> Restart & Run All`.

## Cleaning Up

```python
!rm -rf ~/w3p1_testdir
```

```python

```