# Week 13 Problem 3

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/info490-fa16/blob/master/Week2/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

## Problem 1. Classes and Methods.

In the first part of this week's assignment, we will create a class named `OneFlight` that represents a row in the `2001-1.csv` file (`/home/data_scientist/data/2001/2001-1.csv`).

```python
import numpy as np
import pandas as pd
from nose.tools import assert_equal
```

```python
CSV_PATH = "/home/data_scientist/data/2001/2001-1.csv" # do NOT edit this path
```

We will import only the necessary columns of 2001-1.csv, and drop the missing values with `dropna()`.

```python
columns=["Year", "Month", "DayofMonth", "DepTime", "ArrTime", "Origin", "Dest"]
df = pd.read_csv(CSV_PATH, encoding="latin-1", usecols=columns)
df = df.dropna().reset_index(drop=True)
```

In the following code cell, create a class named `OneFlight` that represents a row of `2001.csv`. Before you begin, be sure to read the rest of this notebook to see the examples and understand the big picture.

In the `OneFlight` class, you should

- Write an `__init__` method that initializes `OneFlight.data` to an empty `pandas.Series`,
- Write a method named `read_line(df, n)` that imports the `n`-th row of the dataframe `df` and stores it in the attribute `OneFlight.data`,
- Write a `__str__` method so that a print statement displays a string representation of `OneFlight`.

See the documentation and the examples below for details.

<!-- grade_id: oneflight_answer, points: 0, grade: false, solution: true -->
```python
class OneFlight(object):
    '''
    Represents a row (one flight) from an airline on-time performance csv file.
    
    Attributes
    ----------
    data: A Pandas Series. If the series is not empty, it has the following labels:
          "Year", "Month", "DayofMonth", "DepTime", "ArrTime", "Origin", "Dest"
         
    Methods
    -------
    __init__: Constructor.
    __str__: A string representation of the class.
    read_line(df, n): Imports the n-th row of a Pandas data frame.
    '''
    
    def __init__(self):
        '''
        Initializes the OneFlight class.
        
        Input parameters
        ----------------
        None
        
        Attributes
        ----------
        OneFlight.data (pandas.Series): set to an empty Series.
        '''
        
        # YOUR CODE HERE
        
    def __str__(self):
        '''
        Returns a string representation of the OneFlight class.
        Your string should have the exact format as the example below.
        
        Example
        -------
        >>> flight = OneFlight()
        >>> print(flight)
        Date: 1/17/2001
        Departed from BWI at 18:06.
        Arrived at CLT at 19:31.
        '''
        
        # YOUR CODE HERE
    
    def read_line(self, df, n):
        '''
        Takes a Pandas data frame and an integer "n", reads the n-th row of the data frame,
        and assigns it to the "data" attribute of the OneFlight class.
        
        Parameters
        ----------
        df: A Pandas.DataFrame.
        n: An integer.
        
        Returns
        -------
        None
        '''
        
        # YOUR CODE HERE
```

We begin by creating a `OneFlight` object.

```python
flight0 = OneFlight()
```

All we have done is create an empty object, and we defined our constructor (i.e., the `init` method) so that the `data` attribute is initialized to an empty `Pandas.Series` object, so the following code cell should print out

```python
>>> print(flight0.data)
```

```
Series([], dtype: float64)
```

```python
print(flight0.data)
```

For the instructor's grades, assignments are autograded. The autograder will use the assert statement to check if the attributes are initialized properly.

<!-- grade_id: flight0_data, points: 3, grade: true, solution: false -->
```python
assert_equal(flight0.data.equals(pd.Series()), True)
```

Now, let's check the `read_line()` method by importing the first row of the `2001.csv` file. We used `pandas.read_csv()` to import `2001.csv` so we should get

```python
>>> flight0.read_line(df, 0)
>>> print(flight0.data)
```

```
Year          2001
Month            1
DayofMonth      17
DepTime       1806
ArrTime       1931
Origin         BWI
Dest           CLT
Name: 0, dtype: object
```

```python
flight0.read_line(df, 0)
print(flight0.data)
```

The autograder will check the assert statements.

<!-- grade_id: flight0_read_line_test, points: 3, grade: true, solution: false -->
```python
assert_equal(flight0.read_line(df, 0), None)

assert_equal(flight0.data["Year"], 2001)
assert_equal(flight0.data["Month"], 1)
assert_equal(flight0.data["DayofMonth"], 17)
assert_equal(flight0.data["DepTime"], 1806)
assert_equal(flight0.data["ArrTime"], 1931)
assert_equal(flight0.data["Origin"], "BWI")
assert_equal(flight0.data["Dest"], "CLT")
```

Let's check one more instance.

```python
flight1 = OneFlight()
```

<!-- grade_id: flight1_data_test, points: 3, grade: true, solution: false -->
```python
assert_equal(flight1.data.equals(pd.Series()), True)
```

```python
flight1.read_line(df, 1000)
print(flight1.data)
```

<!-- grade_id: flight1_tests, points: 3, grade: true, solution: false -->
```python
assert_equal(flight1.read_line(df, 1000), None)

assert_equal(flight1.data["Year"], 2001)
assert_equal(flight1.data["Month"], 1)
assert_equal(flight1.data["DayofMonth"], 2)
assert_equal(flight1.data["DepTime"], 1402)
assert_equal(flight1.data["ArrTime"], 1505)
assert_equal(flight1.data["Origin"], "PIT")
assert_equal(flight1.data["Dest"], "AVP")
```

We also defined a string represention of `OneFlight`, so a simple print statement will print out a detailed information about the flight.

```python
>>> print(flight0)
```

```
Date: 1/17/2001
Departed from BWI at 18:06.
Arrived at CLT at 19:31.
```

```python
>>> print(flight1)
```

```
Date: 1/2/2001
Departed from PIT at 14:02.
Arrived at AVP at 15:05.
```

The print statement will be graded automatically, so you need to reproduce this **exact format**. Note that the string representation of the class is a string of three lines, so you need to use "`\n`" to indicate new lines in the string. Also note that we represent `DepTime` and `ArrTime` in the format of "`hh:mm`", which you can get by applying simple string operations. Instead of combining substrings using "`+`", I suggest you to use [string formatting](https://pyformat.info) so that you can insert variables in a long string. For example, 

```python
a = 9
b = 4
string = '{} is larger than {}.'.format(a, b)
```

makes a long string "`9 is larger than 4.`"

```python
print(flight0)
```

```python
print(flight1)
```

To pass the assertion test, you need to make sure your strings are exactly the same as the answers below. Make sure you do not have extra spaces or missing spaces.  

<!-- grade_id: flight1_str, points: 3, grade: true, solution: false -->
```python
assert_equal(str(flight0), 'Date: 1/17/2001\nDeparted from BWI at 18:06.\nArrived at CLT at 19:31.')
assert_equal(str(flight1), 'Date: 1/2/2001\nDeparted from PIT at 14:02.\nArrived at AVP at 15:05.')
```

## Problem 2. Inheritance.

In the second part of this week's assignment, we will create a child class `AnotherFlight` inherited from `OneFlight`:

```python
class AnotherFlight(OneFlight):
    ...
```

In the `AnotherFlight` class, you should

- Override the `__init__` method so that it accepts an optional argument `data`,
- Write a method named `get_airtime()` that computes the difference between `DepTime` and `ArrTime` in minutes, and
- Write a method named `__lt__` (which stands for "less than") that compares the flight time of two `AnotherFlight` objects.
- *Hint: To avoid repeating codes, you could call `get_airtime()` method inside the `__lt__` method.*

See the documentation and the examples below for details.

<!-- grade_id: anotherflight_answer, points: 0, grade: false, solution: true -->
```python
class AnotherFlight(OneFlight):
    '''
    Inherits from OneFlight.
    Represents a row (one flight) from an airline performance csv file.
    
    Attributes
    ----------
    data: A Pandas Series with the following labels:
          "Year", "Month", "DayofMonth", "DepTime", "ArrTime", "Origin", "Dest"
    
    Methods
    -------
    __init__: Constructor. Optionally accepts a dictionary.
    __lt__: Overrides the < (less than) operator to compare the air times of two AnotherFlight objects.
    get_airtime(): Returns the flight time (the difference betwen "DepTime" and "ArrTime") in minutes.
    '''
    
    def __init__(self, data=None):
        '''
        Initializes the AnotherFlight class.
        
        Input parameters
        ----------------
        data (dict): Optional.
                     If the argument "data" is not None,
                     AnotherFlight.data is built from this dictionary.
        
        Attributes
        ----------
        AnotherFlight.data: A pandas.Series).
                            Set to an empty pandas.Series if there's no argument.
                            If the optional parameter "data" is a dictionary,
                            the constructor uses the dictionary to create a
                            pandas.Series.
        '''
        
        # YOUR CODE HERE
    
    def get_airtime(self):
        '''
        Returns the flight time in minutes, where air time = ArrTime - DepTime.
        For example, if DepTime == 1820 and ArrTime == 1710, get_airtime() returns 70
        because 1820 is 6:20 pm and 1710 is 5:10 pm and the difference is 70 minutes.
        
        Parameters
        ----------
        None
        
        Returns
        -------
        An int.
        '''
        
        # YOUR CODE HERE
    
    def __lt__(self, other):
        '''
        Returns true if the air time of "self" is strictly less than the air time of "other".
        Returns false if the air time of "self" is greater than or equal to "other".
        '''
        
        # YOUR CODE HERE
```

Let's create an `AnotherFlight` object, first without using the optional argument `data`.

```python
flight2 = AnotherFlight()
print(flight2.data)
```

```python
assert_equal(flight2.data.equals(pd.Series()), True)
```

Since `AnotherFlight` inherits from `OneFlight`, the `read_line()` and `__str__` methods of `OneFlight` are also available in its child `AnotherFlight`:

```python
flight2.read_line(df, 0)
print(flight2)
```

```python
assert_equal(flight2.data["Year"], 2001)
assert_equal(flight2.data["Month"], 1)
assert_equal(flight2.data["DayofMonth"], 17)
assert_equal(flight2.data["DepTime"], 1806)
assert_equal(flight2.data["ArrTime"], 1931)
assert_equal(flight2.data["Origin"], "BWI")
assert_equal(flight2.data["Dest"], "CLT")
assert_equal(str(flight2), 'Date: 1/17/2001\nDeparted from BWI at 18:06.\nArrived at CLT at 19:31.')
```

Since the `__init__` method of `AnotherFlight` has been overridden to accept an optional argument `data`, we can also use a dictionary to create an object:

```python
>>> flight3_data = df.iloc[1000].to_dict()
>>> print(type(flight3_data))
```
```
<class 'dict'>
```

```python
>>> print(flight3_data)
```
```
{'Month': 1, 'DepTime': 1402.0, 'Origin': 'PIT', 'Dest': 'AVP', 'ArrTime': 1505.0, 'DayofMonth': 2, 'Year': 2001}
```

```python
>>> flight3 = AnotherFlight(flight3_data)
>>> print(flight3)
```

```
Date: 1/2/2001
Departed from PIT at 14:02.
Arrived at AVP at 15:05.
```

```python
flight3_data = df.iloc[1000].to_dict()
print(type(flight3_data))
print(flight3_data)
```

```python
flight3 = AnotherFlight(flight3_data)
print(flight3)
```

<!-- grade_id: anotherflight_init_test, points: 5, grade: true, solution: false -->
```python
assert_equal(flight3.data["Year"], 2001)
assert_equal(flight3.data["Month"], 1)
assert_equal(flight3.data["DayofMonth"], 2)
assert_equal(flight3.data["DepTime"], 1402)
assert_equal(flight3.data["ArrTime"], 1505)
assert_equal(flight3.data["Origin"], "PIT")
assert_equal(flight3.data["Dest"], "AVP")
assert_equal(str(flight3), 'Date: 1/2/2001\nDeparted from PIT at 14:02.\nArrived at AVP at 15:05.')

# some more tests
flight4 = AnotherFlight({
    'Year': 2001,
    'Month': 1,
    'DayofMonth': 28,
    'DepTime': 1813.0,
    'ArrTime': 1846.0,
    'Origin': 'GSO',
    'Dest': 'CLT'
})

assert_equal(flight4.data["Year"], 2001)
assert_equal(flight4.data["Month"], 1)
assert_equal(flight4.data["DayofMonth"], 28)
assert_equal(flight4.data["DepTime"], 1813)
assert_equal(flight4.data["ArrTime"], 1846)
assert_equal(flight4.data["Origin"], "GSO")
assert_equal(flight4.data["Dest"], "CLT")
assert_equal(str(flight4), 'Date: 1/28/2001\nDeparted from GSO at 18:13.\nArrived at CLT at 18:46.')

flight5 = AnotherFlight({
    'Year': 2001,
    'Month': 1,
    'DayofMonth': 28,
    'DepTime': 1331.0,
    'ArrTime': 1524.0,
    'Origin': 'PHL',
    'Dest': 'DTW'
})

assert_equal(flight5.data["Year"], 2001)
assert_equal(flight5.data["Month"], 1)
assert_equal(flight5.data["DayofMonth"], 28)
assert_equal(flight5.data["DepTime"], 1331)
assert_equal(flight5.data["ArrTime"], 1524)
assert_equal(flight5.data["Origin"], "PHL")
assert_equal(flight5.data["Dest"], "DTW")
assert_equal(str(flight5), 'Date: 1/28/2001\nDeparted from PHL at 13:31.\nArrived at DTW at 15:24.')
```

Now let's test the `get_airtime()` method. `flight2` departed at 18:06 and arrived at 19:31, so the difference is 1 hour and 25 minutes or 85 minutes. `flight3` departed at 14:02 and arrived at 15:05, so the difference is 1 hour and 3 minutes or 63 minutes. 

```python
print(flight2.get_airtime())
print(flight3.get_airtime())
```

<!-- grade_id: get_airtime_test, points: 5, grade: true, solution: false -->
```python
assert_equal(flight2.get_airtime(), 85)
assert_equal(flight3.get_airtime(), 63)

# some more tests
assert_equal(flight4.get_airtime(), 33)
assert_equal(flight5.get_airtime(), 113)
```

Finally, we check the less than operator. 

```python
print(flight2 < flight3)
```

```python
print(flight3 < flight2)
```

<!-- grade_id: less_than_test, points: 5, grade: true, solution: false -->
```python
assert_equal(flight2 < flight3, False)
assert_equal(flight3 < flight2, True) 
assert_equal(flight2 == flight3, False)

# some more tests
assert_equal(flight4 < flight5, True)
assert_equal(flight5 < flight4, False) 
assert_equal(flight4 == flight5, False)
```