# Week 13 Problem 2

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/info490-fa16/blob/master/Week2/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
from nose.tools import assert_equal
```

# Problem 1.

This week, we will be focusing on Object Oriented Programming.  To do this, you will get some more experience with creating a class, as well as creating a few pure functions and modifiers.  So without further ado, let's start by creating a class called "Weather".  This class will have three attributes: temperature (called "temp"), the cumulative precipitation measured in inches (called "in"), and the cumuluative precipitation in feet (called "ft").

A couple things that should be noted as you work through this assignment: 

- You should not have more than 12 inches in 'in' at a time.  Thus, 56 inches should be converted to: ft=4 , in=8
- The temperature can be either in celcius or in fahrenheit
- You MUST properly document everything, ESPECIALLY your class (you must have the docstring body!!)

In the following cell, create your class called Weather.

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

# Problem 2.

Now that you have made your class, lets create a pure function: let's create a function that allows us to see what is contained within a Weather object.  Your function should return output in this format:

(inch, ft, temp)

So, if you have a weather object with 4 ft and 8 inches of precipitation over a year, with a temperature of 56, your function should return the data EXACTLY like this:

(8, 4, 56)

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
def return_Weather(w):
    '''
    Takes a Weather object and returns the data contained within as a tuple
    
    Parameters
    ----------
    w: a weather object
    
    Returns
    -------
    a tuple
    '''
    # YOUR CODE HERE
```

<!-- grade_id: markdown_2, points: 5, grade: true, solution: false -->
```python
urbana = Weather()
urbana.inch = 3
urbana.ft = 4
urbana.temp = 48

assert_equal(return_Weather(urbana), (3, 4, 48))

amazon = Weather()
amazon.inch = 7
amazon.ft = 9
amazon.temp = 87

assert_equal(return_Weather(amazon), (7, 9, 87))
```

# Problem 3.

Now, we need to make a function that will modify a Weather object so that the "inch" attribute is not 12 or more inches, and if there are 12 or more inches, the appropriate number of feet are added to the "ft" attribute.  For example, your function should be able take the following Weather object (presented here by return_Weather).

(15, 4, 56)

And transform it to this:

(3, 5, 56)

Another example of such a transformation is this:
(32, 1, 56) -> (8, 3, 56)

<!-- grade_id: markdown_1, points: 0, grade: false, solution: true -->
```python
def fix_inch(w):
    '''
    Takes a weather object and changes the "inch" attribute so that it does
    not contain 12 or more inches in it, modifying the "ft" attribute in the 
    process.  It then returns the weather object.
    
    Parameters
    ----------
    a Weather object
    
    Returns
    -------
    a Weather object
    '''
# YOUR CODE HERE
```

<!-- grade_id: markdown_3, points: 10, grade: true, solution: false -->
```python
champaign = Weather()
champaign.inch = 40
champaign.ft = 4
champaign.temp = 48
assert_equal(return_Weather(fix_inch(champaign)), (4, 7, 48))

athens = Weather()
athens.inch = 32
athens.ft = 1
athens.temp = 56
assert_equal(return_Weather(fix_inch(athens)), (8, 3, 56))
```

# Problem 4.

Finally, we will create a code that will allow us to change the amount of feet and inches of cumulative precipitation, as well as let us modify the temperature of our Weather object.  Please note that you might want to use fix_inch from above to help you with the formatting of your inches and feet.

<!-- grade_id: add_change, points: 0, grade: false, solution: true -->
```python
def add_dat(w, inch, feet, temp):
    '''
    Adds inches and feet to a weather object, and changes the temperature of the 
    weather object to a new temperature.
    
    Parameters
    ----------
    w: a weather object
    inch: an integer
    feet: an integer
    temp: an integer
    
    Returns
    -------
    A Weather object
    '''
    # YOUR CODE HERE
```

<!-- grade_id: add_change_test, points: 10, grade: true, solution: false -->
```python
assert_equal(return_Weather(add_dat(athens, 1, 0, 57)), (9, 3, 57))
assert_equal(return_Weather(add_dat(urbana, 7, 1, 48)), (10, 5, 48))
assert_equal(return_Weather(add_dat(amazon, 13, 0, 101)), (8, 10, 101))
assert_equal(return_Weather(add_dat(champaign, 6, 0, 0)), (10, 7, 0))
```