```python
%matplotlib inline
```

# Week 13 Problem 1

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/info490-fa16/blob/master/Week2/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
from nose.tools import assert_equal, assert_almost_equal
import numpy as np
import matplotlib.pyplot as plt
```

For this assignment you'll learn how to create a Python class. Most class definitions start with an `__init__` method. This is a special method that is always run when a class is instantiated. One thing you might find peculiar about creating a method within a class is that the `self` argument is always necessary. This special argument allows you to access the class's attributes from within a method. In fact, it is thought of as an instance of the class itself. In addition, the `__init__` method is where you place arguments that the class itself will take. For more reading on classes, methods, and `__init__`, see [this article(https://en.wikibooks.org/wiki/A_Beginner%27s_Python_Tutorial/Classes)

Here is an example of a class called `Robot` with an `__init__` method that takes the obligatory `self` argument as well as a `name` argument.

```python
class Robot:
    
    """
    A class for representing a robot
    
    Parameters
    ----------
    name : string, a name for the Robot
    
    """
    
    def __init__(self, name):
        
        # assign the name to self
        self.name = name
        # create an attribute for the length of the name
        self.letters_in_name = len(name)
        # create private variable for demonstration
        first_letter = name[0]
        
    def print_info(self):
        
        """
        Prints the Robot's name an number of letters
        """
        # create the message and assign as attribute
        self.message = ("This robot's name is {0}. "
                        "It has {1} letters in its name"
                        ).format(self.name, self.letters_in_name)
        # print the message
        print(self.message)
        
```

Notice how by assigning the argument `name` to `self.name`, we were able to access it inside the `print_info` method. Also notice how it is possible to create new attributes from within `__init__` or any other method. First let's intentionally create some errors to demonstrate some key points.

```python
# call Robot() without argument
try:
    Robot()
# print the error
except TypeError as e:
    print(e)
```

The code above tries to call the newly made class and catches and prints the inevitable error. Notice how `Robot`'s `__init__` method only expects to take one argument, `name`. That's because `self` is passed to each method implicitly. In other words, when you are trying to define the arguments your class will take, remember to include self as the first argument.

```python
# get an instance of the class
x = Robot("Andrew")
# try to access an attribute defined locally within __init__
try:
    x.first_letter
# print the error
except AttributeError as e:
    print(e)
```

We defined `first_letter` within `__init__` but it is not accessible as an attribute. Why? Because we never assigned it to `self`! Like any other function, variables defined within it are only accessible from within it. `self` is a vehicle for moving attributes between methods as well as making them accessible from outside the confines of the class.

Try to understand what is happening in the code below. It is critical to understand this in order to complete all the assignments this week and, more importantly, understand Python classes.

```python
# instantiate the class and call the print_info() method
x = Robot('Andrew')
x.print_info()
```

```python
print("The name attribute is            :", x.name)
print("The letters_in_name attribute is :", x.letters_in_name)
print("The message attribute is         :", x.message)
```

```python
# all the attributes in the instance
x.__dict__
```

```python
# now get a different instance of the class. It will have different attributes
y = Robot("Info")
y.print_info()
y.__dict__
```

# Problem 1.

For this problem you will create a simple class called `Circle` that takes one argument, the circle's radius, and creates three attributes:

* `radius`, the radius of the circle
* `diameter`, the diameter of the circle
* `area`, the area of the circle

Note in the unit tests how I create a couple instances of the class and ensure that the attributes are there and return the correct values. You'll want to put all your attribute creation and assignment within an `__init__` method. That is the only method this class needs, but you can create more if you like.

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
#YOUR CODE HERE
```

<!-- grade_id: p1_test, points: 10, grade: true, solution: false -->
```python
# test one instance
a = Circle(1)
assert_equal(a.radius, 1)
assert_equal(a.diameter, 2)
assert_almost_equal(a.area, 3.141592653589793, places=5)
# test another
b = Circle(1.5)
assert_equal(b.radius, 1.5)
assert_equal(b.diameter, 3)
assert_almost_equal(b.area, 7.0685834705770345, places=5)
```

# Problem 2.

Now create another version of your circle class called `Circle2` but this time add a `plot` method that takes two arguments `xlim` and `ylim`, both two-element lists. They will be the x-limits and y-limits of the plot. Center the circle at (0,0).

You may rewrite your `__init__` in your new class, but it would be more Pythonic to [inherit it](https://docs.python.org/2/tutorial/classes.html#inheritance) from your `Circle` class. Here is a rough outline of the `plot` method:

```
def plot( # arguments ):
    
    fig, ax = # the usual function
    
    c = plt.Circle(# different arguments)
    
    ax.add_patch(c)
    
    # set x limit
    # set y limit
    
    return ax
    
```

Since I'm giving you some of the code, I'll leave understanding how to implement it within `Circle2` up to you. This means understanding the format of the `xlim` and `ylim` arguments.

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
#YOUR CODE HERE
```

<!-- grade_id: p2_test, points: 15, grade: true, solution: false -->
```python
# test one case
r = 2; xl=[-3,3]; yl=[-3,3];
c = Circle2(r)
x = c.plot(xl, yl)
d = x.patches[0]
assert_equal(d.center, (0,0))
assert_equal(d.radius, r)
assert_equal(x.get_xlim(), tuple(xl))
assert_equal(x.get_ylim(), tuple(yl))
assert_equal(c.radius, r)
assert_equal(c.diameter, 4)
assert_almost_equal(c.area, 12.566370614359172, places=5)
plt.close()

# test another
r = 4; xl=[-4,4]; yl=[-4,4];
c = Circle2(r)
x = c.plot(xl, yl)
d = x.patches[0]
assert_equal(d.center, (0,0))
assert_equal(d.radius, r)
assert_equal(x.get_xlim(), tuple(xl))
assert_equal(x.get_ylim(), tuple(yl))
assert_equal(c.radius, r)
assert_equal(c.diameter, 8)
assert_almost_equal(c.area, 50.26548245743669, places=5)
plt.close()
```

```python

```

```python

```