# Week 8 Problem 1

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/info490-fa16/blob/master/Week2/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
from nose.tools import assert_equal, assert_true
from numpy.testing import assert_array_equal, assert_almost_equal
import numpy as np
```

# Problem 1.

For this problem you will perform a similar analysis to Week 4 Problem 1.3 where you aim to find the Northernmost and Southernmost airports. This time you will work with slightly different data and will not be limited to the state of Illinois. Write a function that returns the names of the northernmost and southernmost airports in a given state as defined by latitude (higher latitude means farther north). Your function will likely need to perform the following:
1. Read in latitudes from `/home/data_scientist/data/data.csv`
2. Read in airport name and state from `/home/data_scientist/data/data.csv`
3. Find the airports with maximal and minimal latitude for the given state
4. Returns a numpy array of the form `array(['Northernmost', 'Southernmost'])`

You may use any method you with to get the data from csv into NumPy arrays including `np.loadtxt`. If you choose that route, you may run into encoding issues with the strings in the airport and state columns. To get around this, load the data in as `bytes` then convert to `unicode` as such:

`my_array = np.loadtxt('file.csv', dtype=bytes).astype("U")`

Remember that NumPy arrays are homogeneous so you'll have to load numbers and strings in different arrays (or read numbers in as strings). Finally, pay close attention to the output format. You may have the right answer but the wrong dimensions, string encoding, etc. will cause assert tests to fail and you to lose points.

```python
# know thy data!
!head /home/data_scientist/data/data.csv
```

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
#YOUR CODE HERE
```

```python
# Output should be something like (don't worry too much about dtype):
# array(['Waukegan Regional', 'Cairo'], dtype='<U17')
north_south_airports('IL')
```

```python
# Output should be something like (don't worry too much about dtype):
# array(['Princeville', 'Hilo International'], dtype='<U18')
north_south_airports('HI')
```

<!-- grade_id: p1_test, points: 10, grade: true, solution: false -->
```python
# assert tests for California
CA_airpts = north_south_airports("CA")
assert_equal(CA_airpts.shape, (2, ))
CA_airpts_list = [x.strip() for x in CA_airpts]
assert_true('Tulelake' in CA_airpts_list[0])
assert_true('Brown' in CA_airpts_list[1])
# assert tests for Missouri
MO_airpts = north_south_airports("MO")
assert_equal(MO_airpts.shape, (2, ))
MO_airpts_list = [x.strip() for x in MO_airpts]
assert_true('Memphis' in MO_airpts_list[0])
assert_true('Caruthersville' in MO_airpts_list[1])
```

# Problem 2.

The main purpose of this problem is to set up the 3rd problem where you'll perform simple linear regression using numpy. This problem asks you to create a [design matrix](https://en.wikipedia.org/wiki/Design_matrix#Simple_Regression) which is simply an $n$ by 2 matrix where the first column is all ones and the second column is the original array. For example, 

`array([1.,2.,3.])` 

should turn into 

`array([[1., 1.],
       [1., 2.],
       [1., 3.]])`

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
#YOUR CODE HERE
```

<!-- grade_id: p2_test, points: 5, grade: true, solution: false -->
```python
# insert assert tests here
des_mtx = create_design_matrix(np.array([1.1,2.2,3.3]))
assert_array_equal(des_mtx, np.array([[1., 1.1], [1., 2.2], [1., 3.3]]))
```

# Problem 3.

Here you will implement [simple linear regression](https://en.wikipedia.org/wiki/Simple_linear_regression), which is a technique for fitting a linear "trend line" to a set of data points. 
![Simple Linear Regression](https://upload.wikimedia.org/wikipedia/commons/3/3a/Linear_regression.svg)

Let's say we own an ice cream truck and we have data on how many ice cream cones we sold in a day (y-axis) as well as the high temperature for that day (x-axis) over the last year. We wish to use that data to predict the number of ice cream cones we will sell *tomorrow*, when the high temperature will be 50 degrees. We could look back at our data and calculate on average how many ice cream cones we sold on days that were roughly 50 degrees, but that might be tedious and inexact. Instead, we'll fit a regression line (red) to the data (blue) and simply plug in the temperature to predict tomorrow's ice cream cone sales.

Mathematically that looks like this:

$$\hat{y} = \beta_0 + \beta_1x + \epsilon$$

where $\hat{y}$ is the predicted number of ice cream cones sold, $x$ is the high temperature, $\epsilon$ is an error term, and $\beta_0, \beta_1$ are estimated and can be thought of as the *intercept* and *slope*. It is often convenient to write this equation in matrix form as 

$$y = \beta X + \epsilon$$

where $\beta$ is a vector $[\beta_0, \beta_1]$ and $X$ is the design matrix. In order to find $\beta$, you can use the formula

$$\beta = (X^T X)^{-1}(X^T y)$$

In this problem, you will take two 1-d arrays `x` and `y` and estimate the beta coefficients for a simple linear regression. Use numpy's built-in matrix algebra functionality, the `create_design_matrix` function from part 2, and the formula above. To be clear, $X$ is the design matrix, $T$ means transpose, and $-1$ means inverse. That means you need to perform 5 operations:
1. Create the design matrix $X$
1. Compute $X^T X$
2. Invert $X^T X$
3. Compute $X^T y$
4. Return $(X^T X)^{-1}(X^T y)$

<!-- grade_id: p3_answer, points: 0, grade: false, solution: true -->
```python
#YOUR CODE HERE
```

```python
# create x values from -20 to 60
x = np.linspace(-20, 60, 10000)
# create true y values as a function of x and add random noise
y = 5 + 5/22 * x + np.random.randn(len(x))
# create the beta coefficient from the data
beta=slr(x, y)
# print the results
print('Regression Report')
print('-'*30)
print('True beta0 = ', 5)
print("Pred beta0 = ", beta[0])
print('True beta1 = ', 5/22)
print("Pred beta1 = ", beta[1])
print('True ice cream sales for x=50: ', 5 + 5/22 * 50)
print('Pred ice cream sales for x=50: ', np.dot(beta, np.array([1, 50.])))
print('Error:', 5 + 5/22 * 50 - np.dot(beta, np.array([1, 50.])))
```

<!-- grade_id: p3_tests, points: 10, grade: true, solution: false -->
```python
# create x and y with random noise
np.random.seed(1000)
x = np.random.randn(10000)
y = 3.4 + 1.3 * x + np.random.randn(10000)
# compute the beta vector
beta = slr(x, y)
# check that the true beta is close to the estimated beta
assert_almost_equal(np.array([3.4, 1.3]), beta, decimal=1)

# try another b0 and b1
y = 101.4 + 1176.1 * x + np.random.randn(10000)
# compute the beta vector
beta = slr(x, y)
# check that the true beta is close to the estimated beta
assert_almost_equal(np.array([101.4, 1176.1]), beta, decimal=1)
```

```python

```