# Week 12 Problem 1

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/info490-fa16/blob/master/Week2/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
from nose.tools import assert_equal
import pickle, csv, sys, json, os
```

```python
%%bash
mkdir -p ~/w12p1_dir
cd ~/w12p1_dir/
wget -nv https://raw.githubusercontent.com/UI-DataScience/summer2014/c6057185f9a910dfc04dfdc8e65d7f1ef9bb3f64/hw5/%23informatics.p
```

In this assignemnt you will explore various file I/O methods, and to do so you'll be tasked with reading and writing Python dictionaries. This is a good data structure to work with as it does not usually fit into a tabular form easily (this one would, however, because we'll create it from a table). The following function takes no arguments but creates a dictionary of data from `airports.csv` where the keys are the IATA codes and the values are tuples of the rest of the data.

```python
def make_data():
    res = {}
    with open('/home/data_scientist/data/airports.csv', 'r', encoding='latin1') as f:
        for row in csv.reader(f, delimiter=','):
            res[row[0]] = tuple(row[1:])
            
    return res

res = make_data()
```

```python
len(res)
```

```python
res['00M']
```

```python
size_in_mb = sys.getsizeof(res)/1000000
print('res is roughly {0:0.2f} MB'.format(size_in_mb))
```

The following two functions will convert a dictionary to a string and back again. You don't really have to worry about them, but they will be used to check your work for Problems 1 and 2.

```python
def dict_to_str(a_dict):
    
    return str(a_dict)

def str_to_dict(a_dict):
    
    return eval(a_dict)

# verify
str_to_dict(dict_to_str(res)) == res
```

# Problem 1.

Create a funtion that takes a Python string and writes it to a file. This should be a very simple wrapper function that takes two strings, one for the Python object to write out and the other for the file name to write it as.

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
def write_out(obj, fname):
    
    """Wrapper function to write Python objects using standard I/O"""
    
    #YOUR CODE HERE
    
    return None
```

<!-- grade_id: p1_test, points: 4, grade: true, solution: false -->
```python
tdata = 'THIS IS A TEST'
write_out(tdata, '/home/data_scientist/w12p1_dir/p1_test')
ans_p1 = !cat /home/data_scientist/w12p1_dir/p1_test
assert_equal(ans_p1[0], tdata)
assert_equal(len(ans_p1), 1)

tdata_2 = 'THIS IS A\nMULTI-LINE TEST'
write_out(tdata_2, '/home/data_scientist/w12p1_dir/p1_test_2')
ans_p1_2 = !cat /home/data_scientist/w12p1_dir/p1_test_2
assert_equal(ans_p1_2[0], "THIS IS A")
assert_equal(ans_p1_2[1], "MULTI-LINE TEST")
assert_equal(len(ans_p1_2), 2)
```

# Problem 2.

Create a function that is essentially the inverse of Problem 1. That is, it should read in a file using standard Python I/O and create a list with one element per line.

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
def read_in(fname):
    
    """Wrapper function to read standard files"""
    
    #YOUR CODE HERE
    
    return result
```

<!-- grade_id: p2_test, points: 4, grade: true, solution: false -->
```python
# check for a standard case
tdata = "THIS IS A PROBLEM 2 TEST\nIT HAS TWO LINES"
write_out(tdata, '/home/data_scientist/w12p1_dir/p2_test')
p2_ans = read_in('/home/data_scientist/w12p1_dir/p2_test')
assert_equal(len(p2_ans), 2)
assert_equal(p2_ans[0], "THIS IS A PROBLEM 2 TEST\n")
assert_equal(p2_ans[1], "IT HAS TWO LINES")


# now check reading and writing a dictionary!
res = make_data()
res_str = dict_to_str(res)
write_out(res_str, '/home/data_scientist/w12p1_dir/apts.txt')
read_res = read_in('/home/data_scientist/w12p1_dir/apts.txt')
read_res_dict = str_to_dict(read_res[0])
res_two = make_data() # remake data

assert_equal(res, res_two)
assert_equal(read_res_dict, res)
```

# Problem 3.

Now for this problem write a function that takes a Python object and a filename and pickles it with the given filename.

<!-- grade_id: p3_answer, points: 0, grade: false, solution: true -->
```python
def write_pickle(obj, fname):
    
    """Wrapper function to pickle Python objects"""
    
    #YOUR CODE HERE
    
    return None
```

<!-- grade_id: p3_test, points: 4, grade: true, solution: false -->
```python
# insert assert tests here
tdata = make_data()
path = '/home/data_scientist/w12p1_dir/p3.p'
write_pickle(make_data,path)
assert(os.path.exists(path))
# NOTE this only checks if the file is there, not that it's right
```

# Problem 4
Now write a function that takes a filename and returns the un-pickled object at that file location. You can test your functions from problems 3 and 4 by pickling and unpickling random Python objects.

<!-- grade_id: p4_answer, points: 0, grade: false, solution: true -->
```python
def read_pickle(fname):
    
    """Wrapper function to un-pickle Python objects"""
    
    #YOUR CODE HERE
    
    return result
```

<!-- grade_id: p4_test, points: 4, grade: true, solution: false -->
```python
res = make_data()
write_pickle(res, '/home/data_scientist/w12p1_dir/p4.p')
pkl_read_res = read_pickle('/home/data_scientist/w12p1_dir/p4.p')
res_two = make_data()

# you didnt modify the file
assert_equal(res, res_two)
# the pickle loaded file is the same
assert_equal(pkl_read_res, res_two)

# check a different pickle file

pkl_read_res2 = read_pickle("/home/data_scientist/w12p1_dir/#informatics.p")
assert_equal(type(pkl_read_res2), list)
assert_equal(len(pkl_read_res2), 1080)
assert_equal(type(pkl_read_res2[0]), dict)
```

# Problem 5

Now use the same principles as Problems 1-4 to write a Python dictionary as a json. You don't have to worry about handling any object other than a dictionary. This should be extremely straightforward if you use the `json` module and read [the docs](https://docs.python.org/2/library/json.html).

<!-- grade_id: p5_answer, points: 0, grade: false, solution: true -->
```python
def write_json(obj, fname):
    
    """Wrapper function to write dictionaries as JSON"""
    
    #YOUR CODE HERE
    
    return None
```

<!-- grade_id: p5_tests, points: 4, grade: true, solution: false -->
```python
tdata = make_data()
write_json(tdata, '/home/data_scientist/w12p1_dir/p5.json')
ans = !cat /home/data_scientist/w12p1_dir/p5.json
assert_equal(tdata, make_data())
assert('["Jackson Municipal"' in ans[0])
assert_equal(len(ans), 1)
assert_equal(len(ans[0]), 288387)
```

# Problem 6

Finally, write a function that reads the JSON back in as a Python dictionary. You might notice that the resulting dictionary comes back with lists instead of tuples as the entries. Your code should handle this and return a dictionary keyed by airport IATA with tuples as the entries. There are fancy and non-fancy ways to do this, and either is fine. If you want to give yourself a challenge, try some of the suggestions in [this StackExchange article](http://stackoverflow.com/questions/15721363/preserve-python-tuples-with-json).

<!-- grade_id: p6_answer, points: 0, grade: false, solution: true -->
```python
def read_json(fname):
    
    """Read JSON and convert entries of the dictionary to tuples"""
    
    #YOUR CODE HERE
    
    return result
```

<!-- grade_id: p6_tests, points: 5, grade: true, solution: false -->
```python
# insert assert tests here
res = make_data()
write_json(res, '/home/data_scientist/w12p1_dir/p6.json')
json_read_res = read_json('/home/data_scientist/w12p1_dir/p6.json')
res_two = make_data()

assert_equal(res, res_two)
assert_equal(json_read_res, res)
```

## Clean Up

```python
!rm -r /home/data_scientist/w12p1_dir/
```