# Week 11 Problem 1

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/info490-fa16/blob/master/Week2/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
from nose.tools import assert_equal
import numpy as np
import pandas as pd
from numpy.testing import assert_almost_equal, assert_array_equal
```

In this problem, we will use Pandas [apply()](http://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.apply.html) to perform function application in Pandas.

We will use a smaller data set with only 10 rows in this problem, but any function you write in this notebook can be scaled to a larger data set such as `2001.csv`.

```python
df = pd.DataFrame({
    "Origin": ['ORD', 'DFW', 'ATL', 'LAX', 'PHX', 'STL', 'DTW', 'MSP', 'LAS', 'BOS'],
    "Distance": [1050, 370,  228,  678,  873, 472,  370, 1182,  745,  448],
    "ArrDelay": [-3,   4,  23,  10,  20, -3, -10, -12,  -9,  -1],
    "DepDelay": [-4,  -5,  11,  -3,   0, -3,  -8,  -6,   2,   2],
    "Cancelled": [0, 0, 0, 0, 1, 0, 1, 0, 1, 1]
})

print(df)
```

# Problem 1.

Use Pandas [apply()](http://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.apply.html) function to write a function named `get_log10()`, which takes a Pandas data frame as its first argument and a column name (str) as its second argument. It returns a Pandas Series with base-10 log values of the specified column.

Note: [numpy.log()](http://docs.scipy.org/doc/numpy/reference/generated/numpy.log.html) calculates the natural logarithm. For the base-10 logarithm, you should use [numpy.log10()](http://docs.scipy.org/doc/numpy/reference/generated/numpy.log10.html).

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
def get_log10(df, column):
    '''
    Takes a data frame and a column name.
    Returns base-10 log of the specified column.
    
    Parameters
    ----------
    df: A pandas data frame.
    column: A string.
    
    Returns
    -------
    A Pandas Series.
    '''
    
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: p1_test, points: 6, grade: true, solution: false -->
```python
# insert assert tests here
log_10_distance = get_log10(df, "Distance")
answer1 = np.array([
    3.0211893, 2.56820172, 2.35793485, 2.83122969,  2.94101424,
    2.673942, 2.56820172, 3.07261748, 2.87215627, 2.65127801
])

assert_almost_equal(log_10_distance.values, answer1)
```

# Problem 2.

Use Pandas apply() to write a function named `get_sum()` that takes a data frame and returns the sum of each row in the specified columns. That is, if the arrival and departure delay columns of our data frame are:

```
   ArrDelay  DepDelay
0        -3        -4
1         4        -5
2        23        11
3        10        -3
4        20         0
5        -3        -3
6       -10        -8
7       -12        -6
8        -9         2
9        -1         2
```

we want to calculate the sum of each row in these columns:

```
0    -7
1    -1
2    34
3     7
4    20
5    -6
6   -18
7   -18
8    -7
9     1
```

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
def get_sum(df, columns):
    '''
    Takes a Pandas DataFrame as its first argument and a list of columns as its second argument.
    Returns a Pandas Series with the sum of the columns listed in "columns".
    
    Parameters
    ----------
    df: A pandas DataFrame.
    columns: A list of strings.
    
    Returns
    -------
    A pandas Series.
    '''

    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: p2_test, points: 6, grade: true, solution: false -->
```python
# insert assert tests here
total_delay = get_sum(df, ["ArrDelay", "DepDelay"])
answer2 = np.array([ -7,  -1,  34,   7,  20,  -6, -18, -18,  -7,   1])
assert_almost_equal(answer2, total_delay.values)

arr_delay = get_sum(df, ["ArrDelay"])
answer3 = np.array(df.ArrDelay)
assert_almost_equal(answer3, arr_delay.values)
```

# Problem 3

Now we want to apply our own user-defined funtion to a dataframe. This is useful when the `ufuncs` don't concisely do what we need them do. For example, say we wish to create a new series that is the value in the `Origin` column if the flight was cancelled and the string 'NA' otherwise.  Write a function called `get_cancelled_origin` that takes a row of a DataFrame (so a Series) as an argument. Return the 'Origin' element if the 'Cancelled' element equals 1. Otherwise return 'NA'.

Notice we didn't give you the function shell, you must create it yourself. You may use the traditional function definition style or the [lambda function style](http://www.python-course.eu/python3_lambda.php).

<!-- grade_id: p3_answer, points: 0, grade: false, solution: true -->
```python
#YOUR CODE HERE
```

<!-- grade_id: p3_tests, points: 6, grade: true, solution: false -->
```python
s = pd.Series({
        'foo':'bar',
        "Origin":'A', 
        'Cancelled': 1
          })
assert_equal("A", get_cancelled_origin(s))

s = pd.Series({
        'foo':'bar',
        "Origin":'A', 
        'Cancelled': 0
          })

assert_equal("NA", get_cancelled_origin(s))

s = pd.Series({
        'foo':'bar',
        "Origin":'A', 
        'Cancelled': -1
          })

assert_equal("NA", get_cancelled_origin(s))
```

# Problem 4

Now write a function to apply a generic function (such as `get_cancelled_origin`) to a dataframe along a given axis. This shouldn't be too hard because we are not passing any arguments to the generic function.



<!-- grade_id: p4_answer, points: 0, grade: false, solution: true -->
```python
def apply_userdefined_func(df, f, axis=1):
    
    '''
    Generic wrapper to apply a function to a dataframe on a given axis
    
    Parameters
    ----------
    df: A pandas DataFrame.
    f: A function to apply
    axis: the axis argument to pass to pd.DataFrame.apply()
    
    Returns
    -------
    A pandas Series.
    '''
    
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: p4_tests, points: 7, grade: true, solution: false -->
```python
ans4 = pd.Series(
    ["NA", "NA", "NA", "NA", "PHX", "NA", "DTW", "NA", "LAS", "BOS"]
    )
assert_array_equal(ans4, apply_userdefined_func(df, get_cancelled_origin).values)


df2 = df.drop('Origin', axis=1)
f = lambda x: (np.mean(x) + 1 )
ans5 = pd.Series(
    [261.75, 93.25, 66.50, 172.25, 224.50,
    117.5, 89.25, 292.0, 185.75, 113.5]
)
assert_almost_equal(ans5, apply_userdefined_func(df2, f, axis=1).values)
```

We can also use a somewhat complicated lambda-style function to calculate confidence intervals for each column like so:

```python
df2 = df.drop('Origin', axis=1)
f = lambda x: (np.mean(x) - 2 * np.std(x), np.mean(x) + 2 * np.std(x))
apply_userdefined_func(df2, f, axis=0)
```

```python

```