# Week 11 Problem 2

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/info490-fa16/blob/master/Week2/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
from nose.tools import assert_equal, assert_almost_equal
from numpy.testing import assert_array_almost_equal, assert_array_equal
import numpy as np
from functools import reduce
```

Lately, you have been studying on functional programming and it's high time that you got some practice with it! For this problem set, you will be focusing mainly on lambda functions, as well as map(), filter(), and reduce().  You will notice that I have already imported functools for you, so please bear that in mind.  

# Problem 1.

For this first problem, you will be making a function that will take a list of temperatures as well as a string that will convey if the temperatures inside are in Celcius or Fahrenheit.  Using lambda functions, you should manipulate the lists so that the list of temperatures get converted to Kelvin. In order to do this, you should keep in mind the following equations:

If you are converting from Celsius to Kelvin, you simply add 273.15 to the original temperature:
$$Kelvin= Celsius + 273.15 $$

If you are converting from Fahrenheit to Kelvin, the easiest avenue is to first convert the temperature to Celsius, then convert the new Celsius temperatures to Kelvin, as was done above:
$$Celsius = \frac{5}{9}\cdot(Fahrenheit-32) $$
$$Kelvin= Celsius + 273.15 $$

Note: you do not have to worry about c_or_f being anything but "cel" or "far".

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
def to_kelvin(temps, c_or_f):
    '''
    Takes a list and a string that is either "cel" or "far".
    Returns a list of temperatures converted to Kelvin.
    
    Parameters
    ----------
    temps: A list
    c_or_f: A string.
    
    Returns
    -------
    A list.
    '''
    # YOUR CODE HERE
    return new_temp
    
```

<!-- grade_id: cube_test, points: 10, grade: true, solution: false -->
```python
Far = [102, 97.7, 99.14, 100.4, 100.04, -40, 68.5, 32, 0, 75]
assert_array_almost_equal(to_kelvin(Far,'far'), 
                          [312.0388888888889, 309.65, 310.45, 311.15, 310.95, 233.14999999999998, 293.42777777777775, 273.15, 255.3722222222222, 297.0388888888889], 
                          decimal = 3)
assert_equal(len(to_kelvin(Far,'far')), len(Far))

assert_array_almost_equal(to_kelvin([1,2,3,4],'far'),
                         [255.92777777777775, 256.4833333333333, 257.0388888888889, 257.59444444444443],
                         decimal = 3)

Cel = [30.5, 42, 53.78, 2.01, -40.00, -20.56, 0, 32, 14.5]
assert_array_almost_equal(to_kelvin(Cel,'cel'), 
                          [303.65, 315.15, 326.92999999999995, 275.15999999999997, 233.14999999999998, 252.58999999999997, 273.15, 305.15, 287.65], 
                          decimal = 3)
assert_equal(len(to_kelvin(Cel, 'cel')), len(Cel))

assert_array_almost_equal(to_kelvin([1,2,3,4],'cel'),
                         [274.15, 275.15, 276.15, 277.15],
                         decimal = 2)
```

# Problem 2.

In this problem, you will make a function that will take in a list of numbers and will reduce the list down to a sublist that contains numbers that obey one of the following criteria: only even numbers, only numbers that are a multiple of 5, or only numbers that are divisible by 7.  In order to denote these criteria, you will use a function argument called `criteria` that will say what the numbers in the sublist should be divisible by ('two', 'five', 'seven').

For example, say that you started out with a list of numbers: 
$$[0,1,2,3,4,5,6,7,8,9,10]$$

If you were to apply the criteria 'two' then you should recieve the following list:
$$[0,2,4,6,8,10]$$

The criteria 'five' should give:
$$[0,5,10]$$

And finally, the 'seven' criteria should leave you with:
$$[0,7]$$

Note: You do not have to worry about `criteria` being anything but 'two', 'five', or 'seven'.

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
def sublist(num_list, criteria):
    '''
    Takes a list and a criteria that is either 'two', 'five', or 'seven'.
    Returns a list of numbers that match the criteria given.
    
    Parameters
    ----------
    num_list = a list
    criteria = a string
    
    Returns
    -------
    A list.
    '''
    # YOUR CODE HERE
    return result
```

<!-- grade_id: markdown_2, points: 10, grade: true, solution: false -->
```python
assert_array_equal(sublist(range(21), "five"), [0,5,10,15,20])
assert_array_equal(sublist(range(21), "two"), [0,2,4,6,8,10,12,14,16,18,20])
assert_array_equal(sublist(range(21), "seven"), [0,7,14])

fib = [1,1,2,3,5,8,13,21,34,55,89]

assert_array_equal(sublist(fib, "five"), [5,55])
assert_array_equal(sublist(fib, "two"), [2,8,34])
assert_array_equal(sublist(fib, "seven"), [21])

squ = [1,4,9,16,25,36,49]

assert_array_equal(sublist(squ, "five"), [25])
assert_array_equal(sublist(squ, "two"), [4,16,36])
assert_array_equal(sublist(squ, "seven"), [49])
```

# Problem 3.

Finally, you will need to create a function that will find the minimum value in a list of numbers using reduce() and lambda functions.  This should be a one-liner function, so don't overthink it! 

Note: You should NOT be using min() in your function definition below!

<!-- grade_id: markdown_1, points: 0, grade: false, solution: true -->
```python
def sublist(num_list):
    '''
    Takes a list of numbers and returns the smallest argument.
    
    Parameters
    ----------
    num_list = a list
    
    Returns
    -------
    A number.
    '''
    # YOUR CODE HERE
    return result
```

<!-- grade_id: markdown_3, points: 5, grade: true, solution: false -->
```python
assert_equal(sublist(range(5,20)), 5)

fib = [1,1,2,3,5,8,13,21,34,55,89]
assert_equal(sublist(fib), 1)

squ = [1,4,9,16,25,36,49]
assert_equal(sublist(squ),1)

f = np.ndarray.tolist(np.random.randint(-100000, 100000, 25))
assert_equal(sublist(f), min(f))

g = np.ndarray.tolist(np.random.rand(100000))
assert_equal(sublist(g), min(g))

h = [4505, 19358, 38437, 42905, 49030, 38037, 49812, 12049, 3443, 38463, 20812, 34606, 46069, 18470, 21422, 18468, 30372, 37250, 10402, 45221, 7938, 34138, 28968, 25111, 38682, 19162, 22387, 28594, 37128, 257, 25515, 44181, 30400, 10328, 7453, 46862, 48531, 47331, 16164, 12907, 11291, 32066, 49055, 20227, 31078, 22199, 32655, 45027, 34102, 9457, 5421, 45664, 28339, 44039, 30423]
assert_equal(sublist(h), 257)
```

```python

```