# Week 11 Problem 3

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/info490-fa16/blob/master/Week2/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
import re
from nose.tools import assert_equal
```

# Problem 1. Functional programming.

In this problem, we will use functional programming (list comprehensions, map, and filter) to process Twitter data. We will use the same tweets from [Problem 6.2](https://github.com/UI-DataScience/info490-fa16/blob/master/Week6/assignments/Problem_2.ipynb), where you were introduced to list comprehensions but were not required to use functional programming. However, you **must** use functional programming in this problem, because you are only allowed to write your code between `[...]` the square brackets of list comprehensiosn or `list(map(...))` the map function.

```python
tweets = '''
New #job opening at The Ottawa Hospital in #Ottawa - #Clinical #Informatics Specialist #jobs http://t.co/3SlUy11dro
Looking for a #Clinical #Informatics Pharmacist Park Plaza Hospital #jobs http://t.co/4Qw8i6YaJI
Info Session 10/7: MSc in Biomedical Informatics, University of Chicago https://t.co/65G8dJmhdR #HIT #UChicago #informatics #healthcare
Here's THE best #Books I've read on #EHR #HIE #HIPAA and #Health #Informatics http://t.co/meFE0dMSPe
@RMayNurseDir @FNightingaleF Scholars talking passionately about what they believe in. #informatics &amp; #skincare  https://t.co/m8qiUSxk0h
'''.strip().split('\n')

print(tweets)
```

## 1. Use list comprehensions to split the text into words.

Use list comprehensions to split each `tweets` into a list of words. Thus, the `split_into_words()` function will return a nested list of lists. See the test case below for an example.

Note that your answer should go in between the square brackets, and therefore you are forced to use list comprehensions.

<!-- grade_id: split_into_words_answer, points: 0, grade: false, solution: true -->
```python
def split_into_words(tweets):
    '''
    Takes a list of tweets, and returns a nested list of words in all tweets.
    Since words are separated by one or more whitespaces,
    the return value is a list of lists of strings with no whitespace.
    
    Parameters
    ----------
    tweets: a list of strings. Strings have whitespaces.
    
    Returns
    -------
    A nested list of strings. Strings have no whitespace.
    Results from splitting each tweet in tweets by whitespace.
    '''
    
    return [
        # YOUR CODE HERE
    ]
```

Run the following code cell to visually inspect the return value of your function.

```python
words = split_into_words(tweets)
print(words)
```

Make sure that your function passes the following assertion test.

<!-- grade_id: split_into_words_test, points: 5, grade: true, solution: false -->
```python
# assertion test
words_answer = [
    ['New', '#job', 'opening', 'at', 'The', 'Ottawa', 'Hospital', 'in', '#Ottawa', '-',
    '#Clinical', '#Informatics', 'Specialist', '#jobs', 'http://t.co/3SlUy11dro'],
    
    ['Looking', 'for', 'a', '#Clinical', '#Informatics', 'Pharmacist', 'Park', 'Plaza', 'Hospital',
    '#jobs', 'http://t.co/4Qw8i6YaJI'],
    
    ['Info', 'Session', '10/7:', 'MSc', 'in', 'Biomedical', 'Informatics,', 'University', 'of', 'Chicago',
    'https://t.co/65G8dJmhdR', '#HIT', '#UChicago', '#informatics', '#healthcare'],
    
    ["Here's", 'THE', 'best', '#Books', "I've", 'read', 'on', '#EHR', '#HIE', '#HIPAA',
    'and', '#Health', '#Informatics', 'http://t.co/meFE0dMSPe'],
    
    ['@RMayNurseDir', '@FNightingaleF', 'Scholars', 'talking', 'passionately', 'about', 'what', 'they', 'believe', 'in.',
    '#informatics', '&amp;', '#skincare', 'https://t.co/m8qiUSxk0h']
]

assert_equal(words, words_answer)
```

## 2. Use the map function to split the text into words.

Rewrite the `split_into_words()` function using `map()` instead of list comprehensions. `split_into_words()` and `map_into_words()` accept exactly the same arguments and return the same results, except the former uses list comprehensions while the latter uses `map()`.

<!-- grade_id: map_into_words_answer, points: 0, grade: false, solution: true -->
```python
def map_into_words(tweets):
    '''
    Takes a list of tweets, and returns a nested list of words in all tweets.
    Since words are separated by one or more whitespaces,
    the return value is a list of lists of strings with no whitespace.
    
    Parameters
    ----------
    tweets: a list of strings. Strings have whitespaces.
    
    Returns
    -------
    A nested list of strings. Strings have no whitespace.
    Results from splitting each tweet in tweets by whitespace.
    '''
    
    return list(map(
        # YOUR CODE HERE
    ))
```

Run the following code cell to visually inspect the return value of your function.

```python
mapped_words = map_into_words(tweets)
print(mapped_words)
```

Make sure that your function passes the following test case. Note that the answer is the same as the one for `split_into_words()`.

<!-- grade_id: map_into_words_test, points: 5, grade: true, solution: false -->
```python
# assertion test
assert_equal(mapped_words, words_answer)
```

## Flatten the nested list (code provided).

It would be useful to flatten the nested list and have a flat list of strings. I'm not going to ask you to do this because a double for loop in list comprehensions can be a little confusing to someone new to list comprehensions. In keeping with the theme of functional programming, here's a function that flattens a list:

```python
def flatten(tweets):
    '''
    Flattens a nested list.
    
    Parameters
    ----------
    tweets: a list of lists.
    
    Returns
    -------
    A flat list.
    '''

    return [
        t for tweet in tweets for t in tweet
    ]
```

The outer loop comes first and the inner loop later in list comprehensions. This is equivalent to

```python
a_list = []
for tweet in tweets:
    for t in tweet:
        a_list.append(t)
return a_list
```

Let's check if our function does what it's supposed to do.

```python
# check the flatten() function
flattened_words = flatten(words)

flattened_words_answer = [
    'New', '#job', 'opening', 'at', 'The', 'Ottawa', 'Hospital', 'in', '#Ottawa', '-',
    '#Clinical', '#Informatics', 'Specialist', '#jobs', 'http://t.co/3SlUy11dro',
    
    'Looking', 'for', 'a', '#Clinical', '#Informatics', 'Pharmacist', 'Park', 'Plaza', 'Hospital',
    '#jobs', 'http://t.co/4Qw8i6YaJI',
    
    'Info', 'Session', '10/7:', 'MSc', 'in', 'Biomedical', 'Informatics,', 'University', 'of', 'Chicago',
    'https://t.co/65G8dJmhdR', '#HIT', '#UChicago', '#informatics', '#healthcare',
    
    "Here's", 'THE', 'best', '#Books', "I've", 'read', 'on', '#EHR', '#HIE', '#HIPAA',
    'and', '#Health', '#Informatics', 'http://t.co/meFE0dMSPe',
    
    '@RMayNurseDir', '@FNightingaleF', 'Scholars', 'talking', 'passionately', 'about', 'what', 'they', 'believe', 'in.',
    '#informatics', '&amp;', '#skincare', 'https://t.co/m8qiUSxk0h'
]

assert_equal(flattened_words, flattened_words_answer)
```

## 3. Use list comprehensions to find long words.

Use list comprehensions to write a function named `get_long_words()` that takes a list of strings. It finds every word that has 10 characters or more and returns a list of tuples `(str, int)`, where the string is a word with 10 or more characters and the integer is the length of that word. See the test case below for what the return value looks like.

<!-- grade_id: get_long_words_answer, points: 0, grade: false, solution: true -->
```python
def get_long_words(words):
    '''
    Takes a list of strings. Finds every word with 10 or more characters.
    Returns a list of tuples of the form
    (a word with 10 or more characters, the length of that word)
    
    Parameters
    ----------
    words: a list of strings.
    
    Returns
    -------
    A list of tuples (str, int).
    '''

    return [
        # YOUR CODE HERE
    ]
```

```python
long_words = get_long_words(flattened_words)
print(long_words)
```

<!-- grade_id: get_long_words_test, points: 5, grade: true, solution: false -->
```python
# assertion test
long_words_answer = [
    ('#Informatics', 12),
    ('Specialist', 10),
    ('http://t.co/3SlUy11dro', 22),
    ('#Informatics', 12),
    ('Pharmacist', 10),
    ('http://t.co/4Qw8i6YaJI', 22),
    ('Biomedical', 10),
    ('Informatics,', 12),
    ('University', 10),
    ('https://t.co/65G8dJmhdR', 23),
    ('#informatics', 12),
    ('#healthcare', 11),
    ('#Informatics', 12),
    ('http://t.co/meFE0dMSPe', 22),
    ('@RMayNurseDir', 13),
    ('@FNightingaleF', 14),
    ('passionately', 12),
    ('#informatics', 12),
    ('https://t.co/m8qiUSxk0h', 23)
]

assert_equal(len(long_words), len(long_words_answer))
assert_equal(set(long_words), set(long_words_answer))
```

## 4. Rewrite get\_long\_words() using map() and filter().

Use `map()` and `filter()` to rewrite `get_long_words()`. `get_long_words()` and `map_long_words()` accept exactly the same arguments and return the same results, except the former uses list comprehensions while the latter uses `map()` and `filter()`.

<!-- grade_id: map_long_words_answer, points: 0, grade: false, solution: true -->
```python
def map_long_words(words):
    '''
    Takes a list of strings. Finds every word with 10 or more characters.
    Returns a list of tuples of the form
    (a word with 10 or more characters, the length of that word)
    
    Parameters
    ----------
    words: a list of strings.
    
    Returns
    -------
    A list of tuples (str, int).
    '''
    
    return list(map(
        # YOUR CODE HERE
    ))
```

```python
mapped_long_words = map_long_words(flattened_words)
print(mapped_long_words)
```

<!-- grade_id: map_long_words_test, points: 5, grade: true, solution: false -->
```python
# assertion test
assert_equal(len(mapped_long_words), len(long_words_answer))
assert_equal(set(mapped_long_words), set(long_words_answer))
```

## 5. Find hashtags using list comprehensions.

Use list comprehensions to find all hashtags (words that start with `#`).

<!-- grade_id: get_hashtags_answer, points: 0, grade: false, solution: true -->
```python
def get_hashtags(words):
    '''
    Takes a list of strings.
    Returns a list of strings that all start with #.
    
    Parameters
    ----------
    words: a list of strings.
    
    Returns
    -------
    A list of strings.
    '''
    
    return [
        # YOUR CODE HERE
    ]
```

```python
hashtags = get_hashtags(flattened_words)
print(hashtags)
```

<!-- grade_id: get_hashtags_test, points: 5, grade: true, solution: false -->
```python
# assertion test
hashtags_answer = [
    '#job', '#Ottawa', '#Clinical', '#Informatics', '#jobs',
    '#Clinical', '#Informatics', '#jobs', '#HIT', '#UChicago',
    '#informatics', '#healthcare', '#Books', '#EHR', '#HIE',
    '#HIPAA', '#Health', '#Informatics', '#informatics', '#skincare'
]

assert_equal(len(hashtags), len(hashtags_answer))
assert_equal(set(hashtags), set(hashtags_answer))
```

## 6. Find hashtags using map() and filter().

Rewrite `get_hashtags()` using `map()` and `filter()`. `get_hashtags()` and `map_hashtags()` accept exactly the same arguments and return the same results, except the former uses list comprehensions while the latter uses `map()` and `filter()`.

<!-- grade_id: map_hashtags_answer, points: 0, grade: false, solution: true -->
```python
def map_hashtags(words):
    '''
    Takes a list of strings.
    Returns a list of strings that all start with #.
    
    Parameters
    ----------
    words: a list of strings.
    
    Returns
    -------
    A list of strings.
    '''
    
    return list(map(
        # YOUR CODE HERE
    ))
```

```python
mapped_hashtags = map_hashtags(flattened_words)
print(mapped_hashtags)
```

<!-- grade_id: map_hashtags_test, points: 5, grade: true, solution: false -->
```python
# assertion test
assert_equal(len(mapped_hashtags), len(hashtags_answer))
assert_equal(set(mapped_hashtags), set(hashtags_answer))
```