# Week 2 Assignment
## Problem 2

The purpose of this problem is the following:
- Run an IPython notebook
- Understand how assignments work
- Understand assignment statements
- Understand how to assign values to variables

Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select `Kernel` → `Restart`) and then run all cells (in the menubar, select `Cell` → `Run All`).

Make sure to fill in any place that says `YOUR CODE HERE`.

Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`.

In the following code cell, assign a value of 365 to an integer variable named `the_integer`.

<!-- grade_id: int, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

In the following code cell, you will see two assert statements, and almost
every assignment will have such tests. These are what are called unit tests.
When you submit completed assignments, instructors will determine your assignment grades by using automated grading tools to perform unit testing on your code, so if you pass all the tests without an error, you get a perfect score.

It should be obvious to you that these tests are valuable hints! Study them. Write each piece of code so that it doesn't produce any errors. Writing the tests first, and then writing the actual code is actually a perfectly valid approach (preferred by many) to programming. It's called test-driven development.

So, remember to learn what assert means, and study nose if you want to get a good grade.

<!-- grade_id: int-test, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(the_integer,365)
assert_equal(type(the_integer), int)
```

In the following code cell, assign the value 1.6180339 to a floating point variable named `floating_bubbles`.

<!-- grade_id: float, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: float-test, points: 5, grade: true, solution: false -->
```python
assert_equal(floating_bubbles, 1.6180339)
assert_equal(type(floating_bubbles), float)
```

Finally, you will need to assign the string 'Spamelot' to a string variable named `the_string`.

<!-- grade_id: string, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: string-test, points: 10, grade: true, solution: false -->
```python
assert_equal(the_string, 'Spamelot')
assert_equal(type(the_string), str)
```

Before you turn this in, please make sure that there are no error messages, otherwise you might lose some points on this assignment.  Good luck!