# Week 2 Problem 2

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.



```python
from nose.tools import assert_equal
```

## Problem 1

The main purpose of this problem is to learn how to assign values to a variable.

In the following code cell, assign a value of 365 to an integer variable named `the_integer`.

<!-- grade_id: int, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

In the following code cell, you will see two assert statements, and almost
every assignment will have such tests. These are what are called unit tests.
When you submit completed assignments, instructors will determine your assignment grades by using automated grading tools to perform unit testing on your code, so if you pass all the tests without an error, you get a perfect score.

It should be obvious to you that these tests are valuable hints! Study them. Write each piece of code so that it doesn't produce any errors. Writing the tests first, and then writing the actual code is actually a perfectly valid approach (preferred by many) to programming. It's called test-driven development.

So, remember to learn what assert means, and study nose if you want to get a good grade.

<!-- grade_id: int-test, points: 5, grade: true, solution: false -->
```python
assert_equal(the_integer,365)
assert_equal(type(the_integer), int)
```

In the following code cell, assign the value 1.6180339 to a floating point variable named `floating_bubbles`.

<!-- grade_id: float, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: float-test, points: 5, grade: true, solution: false -->
```python
assert_equal(floating_bubbles, 1.6180339)
assert_equal(type(floating_bubbles), float)
```

Finally, you will need to assign the string 'Spamelot' to a string variable named `the_string`.

<!-- grade_id: string, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: string-test, points: 10, grade: true, solution: false -->
```python
assert_equal(the_string, 'Spamelot')
assert_equal(type(the_string), str)
```

Before you turn this in, please make sure that there are no error messages, otherwise you might lose some points on this assignment.  Good luck!