# Week 2 Problem 3

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

***

The purpose of this problem is the following:
- Do basic math operations in Python
- Practice Python built-in functions
- Continue to practice assign values to variables
- Learn to use the math library

```python
from nose.tools import assert_equal
import math
```

In the following code cell, calculate the square root of the integer 490490. Assign your result to a floating point variable named `square_root`.

```python
# YOUR CODE HERE
```

<!-- grade_id: w2p3_1, points: 5, grade: true, solution: false -->
```python
assert_equal(type(square_root), float)
assert_equal(square_root, 700.3499125437227)
```

In the following code cell, assign the value 4.9 to a floating point variable named `sphere_radius`.

```python
# YOUR CODE HERE
```

<!-- grade_id: w2p3_2, points: 5, grade: true, solution: false -->
```python
assert_equal(type(sphere_radius), float)
assert_equal(sphere_radius, 4.9)
```

The sphere volume is given by the formula: V = $\frac{4}{3}$ $\pi$ $r^{3}$. In the following cell, calculate the volume of a sphere with a radius equal to `sphere_radius`. Assign your result to a floating point variable named `sphere_volume`.

```python
# YOUR CODE HERE
```

<!-- grade_id: w2p3_3, points: 5, grade: true, solution: false -->
```python
assert_equal(type(sphere_volume), float)
assert_equal(sphere_volume, 492.80697880291416)
```

Finally, in the following cell, change the float varialbe `sphere_volume` to an integer first, and then change that integer to a string. Assign the resulting string to a string variable named `the_volume_string`.

```python
# YOUR CODE HERE
```

<!-- grade_id: w2p3_4, points: 5, grade: true, solution: false -->
```python
assert_equal(type(the_volume_string), string)
assert_equal(the_volume_string, '492')
```