# Week X Problem 1

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/info490-fa16/blob/master/Week2/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
from nose.tools import assert_equal
```

# Problem 1.

The main purpose of this problem is...

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cube_test, points: 10, grade: true, solution: false -->
```python
# insert assert tests here
```

# Problem 2.

The main purpose of this problem is...

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: markdown_2, points: 10, grade: true, solution: false -->
```python
# insert assert tests here
```

# Problem 3.

The main purpose of this problem is...

<!-- grade_id: markdown_1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: markdown_3, points: 10, grade: true, solution: false -->
```python
# insert assert tests here
```