# Week 5 Problem 1

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/info490-fa16/blob/master/Week2/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

***

**Be very careful before you execute the commands you write. You could inadvertently overwrite something you don't intend to.** The reason we create a brand new directory in the beginning of each Unix-based assignment is to give you a "sandbox" or "playground" where you can make mistakes without causing your file system undue harm. It is important that every command you execute is operating on and within the directories we specify. This may sound obvious, by it causes you (and us) undue stress when a command deletes your assignment or home directory, for example.

This assignment is based on chapters 17 and 18 of [TLCL](http://sourceforge.net/projects/linuxcommand/?source=dlp)

```python
from nose.tools import assert_equal
```

Let's make our sandbox directory `w5p1_testdir` as well as 100 subdirectories within it. Within those 100 subdirectories we create 78 empty files to operate on. The purpose of this is to give you a directory strucutre to operate on that is too complex and large to piece through by hand, much like one you'd see in the real world.

```python
%%bash
# make directories 1-100
mkdir -p ~/w5p1_testdir/dir-{001..100}
# make files A-Z as txt, csv, and markdown in each
touch ~/w5p1_testdir/dir-{001..100}/file-{A..Z}.{txt,csv,md}
```

The following cell just check some things to make sure everything is in place as it should be. 

```python
#is w5p1_testdir in ~?
is_in_home = !ls ~ | grep w5p1_testdir | wc -l
if is_in_home[0] != '1':
    raise Exception(
        "w5p1_testdir isn't in /home/data_scientist."\
        " Did you run the cell above?")

# check the number of directories
num_dirs = !ls ~/w5p1_testdir | wc -l

# if it's not what we expect, complain
if num_dirs[0] != "100":
    raise Exception(
        "There should be {0} total directories, we counted {1}."\
        " Did you run the last cell?".format(100, num_dirs[0]))


# set up our expected file count
num_dirs = 100
files_per_dir = 26 * 3
total_files = num_dirs * files_per_dir

# actually count the regular files within w5p1_testdir
files_counted = !find ~/w5p1_testdir/* -type f | wc -l

# if it's not what we expect, complain
if files_counted[0] != "7800":
    raise Exception(
        "There should be {0} total files, we counted {1}."\
        " Did you run the last cell?".format(total_files, files_counted[0]))
```

Notice that by using `find ~/w5p1_testdir/* -type f` I can list all the regular files within `~/w5p1_testdir` **and** its subdirectories.

# Problem 1.

The purpose of this problem is to give you experience pulling files from the internet. Write a `bash` script that 
1. Makes a directory in ~ called "Downloads" if it doesnt exist already (-p flag)
2. Gets the Week 5 Lesson 1 Markdown from the course GitHub repository
3. Outputs the markdown file as ~/Downloads/foo.md

Hint: To pull raw files from GitHub you should click the "Raw" button in the upper right corner. See [here](http://stackoverflow.com/questions/4604663/download-single-files-from-github)

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
%%bash
#YOUR CODE HERE
```

<!-- grade_id: p1_test, points: 5, grade: true, solution: false -->
```python
# check if Downloads exists and where
downloads_dir = !find ~/ -type d -name "Downloads"
# check if foo.md exists and where
foo_file = !find ~/Downloads -type f -name "foo.md"
# make sure its the correct markdown
is_1 = !grep -w "^# Week 5 Lesson 1" ~/Downloads/foo.md | wc -l

assert_equal(downloads_dir[0], '/home/data_scientist/Downloads')
assert_equal(foo_file[0], '/home/data_scientist/Downloads/foo.md')
assert_equal(is_1[0], '1')
```

# Problem 2.

The purpose of this problem is to give you exposure to more advanced commands for manipulating the file system. Copy the `foo.md` file into every directory inside `~/w5p1_testdir`, but **don't copy the file into `~/w5p1_testdir` itself**. If you must, you can simply remove foo.md from `~/w5p1_testdir`, but try to do it without using a `rm` command. The "recommended" ways involve `find` or `xargs`.

Recommended Google Search: *Copy a file into multiple folders* or *copy a file into multiple directories*.

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
%%bash
#YOUR CODE HERE
```

<!-- grade_id: p2_test, points: 10, grade: true, solution: false -->
```python
number_dirs = !ls /home/data_scientist/w5p1_testdir/ | wc -l
number_foo = !find ~/w5p1_testdir/* -type f | grep foo.md | wc -l
assert_equal(int(number_dirs[0]), 100)
assert_equal(int(number_foo[0]), 100)
```

# Problem 3.

The main purpose of this problem is to give you exposure with compressing and archiving files. Create an archive called `odd_dirs.tar` within `~/w5p1_testdir` (so the full filepath is `~/w5p1_testdir/odd_dirs.tar` that contains the **odd-numbered** directories within `~/w5p1_testdir`. Note the following:

* Use the `tar` command with the `cf` option
* The most challenging part will likely be creating a wildcard to grab only the odd-numbered directories
* Start with an `echo` command to try to print out just the odd numbered directories to get your wildcard correct
* After that just add it to a standard tar command with the `--wildcards` flag

In the end you should have 50 directories, each with 79 files in them. Since `tar tf` will count the directories themselves as files, there should be 80 * 50 = 4000 items total in the tarball.

<!-- grade_id: p3_answer, points: 0, grade: false, solution: true -->
```python
%%bash
#YOUR CODE HERE
```

<!-- grade_id: p3_test, points: 10, grade: true, solution: false -->
```python
files_in_tarball = !tar tf ~/w5p1_testdir/odd_dirs.tar | wc -l
foos_in_tarball = !tar tf ~/w5p1_testdir/odd_dirs.tar | grep foo | wc -l
dirs_in_tarball = !tar tf ~/w5p1_testdir/odd_dirs.tar | grep "/$" | wc -l

assert_equal(int(dirs_in_tarball[0]), 50)
assert_equal(int(foos_in_tarball[0]), 50)
assert_equal(int(files_in_tarball[0]), 4000)
```

## Cleaning Up
Also note that this cell can be run to "start over". If you think your code is correct but don't pass the assert tests, run the following and rerun your code. You may be using this alot!

```python
!rm -rf /home/data_scientist/w5p1_testdir/
!rm -rf /home/data_scientist/Downloads/
```