# Week 5 Problem 2

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/info490-fa16/blob/master/Week2/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
from nose.tools import assert_equal
import os
```

# Problem 1.

In this problem set, we will be practicing opening, reading, and then closing a csv file.  For this first section, you will need to write a function that extracts the n-th column of the csv files that are made below.

So here is some important information for this problem:
- Write a function named get_column(filename, n, header=True) that reads the n-th column from a file and returns a list of integers.
- You may assume that the column is made of integers.
- We will also use the optional argument header because the first line of our file lists the names of the columns, but we might want to turn this off to handle a file that doesn't have a header.
- Use a combination of with statement and open() function to open filename in the get_column() function.
- When I tried to use open() to read 2001.csv, I had the following error:
    'utf-8' codec can't decode byte 0xe4 in position 343: invalid continuation byte
- You can avoid this error by using encoding="latin-1" option in open().
- Skip the first line if the header parameter is True; do not skip if it's False.
- Some columns have missing values "NA", and you need a way to handle these missing values. If the n-th column is missing, you should not include that column in result; that is, skip all rows with "NA. As a result, lists returned from different columns may have different lengths.


<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
def get_column(filename, n, header=True):
    '''
    Returns a list from reading the specified column in the CSV file.

    Parameters
    __________
    filename (str): Input file name. The file is assumed to be in a 
                    Comma Separated Values (CSV) format with latin-1 encoding.
    n (int): Column number. The first column starts at 0. The column must be
             a list of integers.
    header (bool): If True, the first line of file is column names.
                   Default: True.
    '''
    result = []

    # YOUR CODE HERE

    return result
```

<!-- grade_id: cube_test, points: 10, grade: true, solution: false -->
```python
csv_with_header = '''
Year,Month,DayofMonth,DayOfWeek,DepTime,CRSDepTime,ArrTime,CRSArrTime,UniqueCarrier,FlightNum,TailNum,ActualElapsedTime,CRSElapsedTime,AirTime,ArrDelay,DepDelay,Origin,Dest,Distance,TaxiIn,TaxiOut,Cancelled,CancellationCode,Diverted,CarrierDelay,WeatherDelay,NASDelay,SecurityDelay,LateAircraftDelay
2001,1,17,3,1806,1810,1931,1934,US,375,N700äæ,85,84,60,-3,-4,BWI,CLT,361,5,20,0,NA,0,NA,NA,NA,NA,1
2001,1,18,4,1805,1810,1938,1934,US,375,N713äæ,93,84,64,4,-5,BWI,CLT,361,9,20,0,NA,0,NA,NA,NA,NA,1
2001,1,19,5,1821,1810,1957,1934,US,375,N702äæ,96,84,80,23,11,BWI,CLT,361,6,10,0,NA,0,NA,NA,NA,NA,NA
2001,1,20,6,1807,1810,1944,1934,US,375,N701äæ,97,84,66,10,-3,BWI,CLT,361,4,27,0,NA,0,NA,NA,NA,NA,NA
'''.strip().encode('latin-1')

csv_no_header = '''
2001,1,17,3,1806,1810,1931,1934,US,375,N700äæ,85,84,60,-3,-4,BWI,CLT,361,5,20,0,NA,0,NA,NA,NA,NA,1
2001,1,18,4,1805,1810,1938,1934,US,375,N713äæ,93,84,64,4,-5,BWI,CLT,361,9,20,0,NA,0,NA,NA,NA,NA,1
2001,1,19,5,1821,1810,1957,1934,US,375,N702äæ,96,84,80,23,11,BWI,CLT,361,6,10,0,NA,0,NA,NA,NA,NA,NA
2001,1,20,6,1807,1810,1944,1934,US,375,N701äæ,97,84,66,10,-3,BWI,CLT,361,4,27,0,NA,0,NA,NA,NA,NA,NA
2001,1,21,7,1810,1810,1954,1934,US,375,N768äæ,104,84,62,20,0,BWI,CLT,361,4,38,0,NA,0,NA,NA,NA,NA,1
'''.strip().encode('latin-1')

with open('test.header.csv', 'wb') as f:
    f.write(csv_with_header)
    
with open('test.noheader.csv', 'wb') as f:
    f.write(csv_no_header)

# header cases
assert_equal(get_column('test.header.csv', 0, header=True), [2001] * 4)
```

First we test the header==True cases:

<!-- grade_id: markdown_2, points: 10, grade: true, solution: false -->
```python
assert_equal(get_column('test.header.csv', 4, header=True), [1806, 1805, 1821, 1807])
assert_equal(get_column('test.header.csv', 28, header=True), [1] * 2)
assert_equal(get_column('test.header.csv', 27, header=True), [])
assert_equal(get_column('test.header.csv', 2, header=True), list(range(17, 21)))
```

Finally, we test the header==False cases:

<!-- grade_id: markdown_3, points: 10, grade: true, solution: false -->
```python
assert_equal(get_column('test.noheader.csv', 0, header=False), [2001] * 5)
assert_equal(get_column('test.noheader.csv', 2, header=False), list(range(17, 22)))
assert_equal(get_column('test.noheader.csv', 4, header=False), [1806, 1805, 1821, 1807, 1810])
assert_equal(get_column('test.noheader.csv', 27, header=False), [])
assert_equal(get_column('test.noheader.csv', 28, header=False), [1] * 3)
```

```python
# Remove the test csv files:
# clean up
os.remove('test.header.csv')
os.remove('test.noheader.csv')
```