# Week 4 Problem 2

## Summary

In this problem set, you will be working with strings and learning how to manipulate them.

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/accy570-fa16/blob/master/Week1/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
from nose.tools import assert_equal
```

# Problem 1.

For this problem, you will be making a function definition that slices a string called `slice_string`.  This function will recieve a string, and two indices.  The string will then be sliced at the provided indices.  For example, if the string was `'animal'` and the indices provided are `0` and `3`, then the function should return the string `'ani'`.  If the indices are the same or if the first is larger than the second, then the `slice_string` function will return an empty string: `''`.

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
def slice_string(a_string, ind1, ind2):
    # YOUR CODE HERE
```

<!-- grade_id: cube_test, points: 5, grade: true, solution: false -->
```python
assert_equal(slice_string('animal', 0, 3),'ani')
assert_equal(slice_string('troll',4,3),'')
assert_equal(slice_string('dessert', 4, 6),'er')
assert_equal(slice_string('Python',2,-2),'th')
```

# Problem 2.

In this part of the problem, we will create a function definition that checks to see if a letter or substring is in the string provided.  If the letter is in the string, then `True` is returned. If the letter is NOT found in the string, `False` must be returned. 

Similarly, if the substring is found in order in the string, then `True` will be returned.  However, if the substring is not matched perfectly with a section of the string, `False` will be returned; for example, if you search for `lon` inside of `lion`, the function should return `False`

This function will be case sensitive; for example, if you search for the letter `'m'` within `'Monty'`, `False` should be returned.

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
def is_in(a_string, letter):
    # YOUR CODE HERE
```

<!-- grade_id: markdown_2, points: 5, grade: true, solution: false -->
```python
assert_equal(is_in('monty','m'), True)
assert_equal(is_in('Monty','m'), False)
assert_equal(is_in('spam','a'), True)
assert_equal(is_in('locket', 'ckt'), False)
```

# Problem 3.

In this final function, it will take a string, a letter, and two indices and will do three things:
- It will see if a letter is located within the original string.
- It will slice up the string at the provided indices
- It will then output a string that is of the following form: "The string is \_ . Is the letter \_ inside the string? \_. The slice of the string is \_.

For example, the output of this:

`list_output('Pickle', 'i', 2, 5)`

Should be this:

`'The string is Pickle. Is the letter i inside the string? Yes. The slice of the string is ckl.'`

Another example coule be: 

`list_output('crunch', 'y', 0, 5)`

Should be this:

`'The string is crunch. Is the letter y inside the string? No. The slice of the string is crunc.'`

<!-- grade_id: markdown_1, points: 0, grade: false, solution: true -->
```python
def list_output(a_string, letter, ind1, ind2):
    # YOUR CODE HERE
```

<!-- grade_id: markdown_3, points: 10, grade: true, solution: false -->
```python
assert_equal(type(list_output('Pickle', 'i', 2, 5)), str)
assert_equal(list_output('moon', 'i', 0, 3),'The string is moon. Is the letter i inside the string? No. The slice of the string is moo.')
assert_equal(list_output('galaxy', 'y', 1, 3), 'The string is galaxy. Is the letter y inside the string? Yes. The slice of the string is al.')
assert_equal(list_output('runner', 'R', 0, 5), 'The string is runner. Is the letter R inside the string? No. The slice of the string is runne.')
```

```python

```