# Week 4 Problem 1

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/accy570-fa16/blob/master/Week1/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.
6. **MAKE SURE YOU SUBMIT YOUR ASSIGNMENT BY CLICKING SUBMIT!**

**Hint:** You might find [this notebook](https://github.com/ProfessorBrunner/rp-pdm15/blob/master/Week1/unixde.ipynb) helpful for some of these problems. I think it is a good supplement to the notebook that was part of your [required reading](https://github.com/UI-DataScience/info490-fa16/blob/master/Week4/notebooks/unixdp.ipynb). 

```python
%%bash

# set a name for the test directory
dir='w4p1_dir'

# make a directory for testing
# remove it and remake if it already exists
if [[ -e ~/$dir ]]; then
    rm -rf ~/$dir 
fi

mkdir ~/$dir
```

```python
from nose.tools import assert_equal
```

# Problem 1.

The main purpose of this problem is to give you exposure to basic command-line data processing. This is useful when files are large and difficult to move or read into another platfrom.

Your task is to count the number of airports in Illinois using the `~/data/airports.csv` file. This file has one entry per airport, so really what you want to know is how many times "IL" appears in the 4th column. Here's what the first few rows look like:

```python
!head ~/data/airports.csv
```

*** 
## Hint

One possible way to do this is to pull out the 4th column and count the occurences of "IL" using `cut`, `grep`, and `wc`. Another way is to use `awk`. Keep in mind that the string entries of the `csv` file actually have quotes around them as shown above. This means that if you search the entire line, you'll have to look for `"IL"` instead of `IL`. You may have to *escape* the `"` characters with a backslash because `awk` and `grep` will treat them as special characters by default. This may agitate the syntax highlighting within the Jupyter Notebook, but don't worry about it. 

Here's an example of what I mean. Pretend we have a file with three lines, `fox`, `firefox`, and `"fox"` and want to return just the line with `"fox"`.:

```python
# make a text file
!echo "fox\n"firefox"\n\"fox\"" > temp.txt
!echo temp file created as temp.txt
```

```python
# this is what the file looks like. it has 3 lines
!cat temp.txt
```

```python
# look for fox but they're all found, including "fox"
!grep fox temp.txt
```

```python
# look for "fox" but they're all found again
!grep "fox" temp.txt
```

```python
# we need to look specifically for "fox" by telling grep we really want the " chars
!grep \"fox\" temp.txt
```

```python
!rm temp.txt
!echo temp.txt removed
```

If you simply use `grep "IL"`, you'll find lines that have `IL` in them as well as lines that have `"IL"` in them, and you only want the latter. Finally, see the notebook linked above for examples of using `awk`.

***

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
num_apts_in_IL = !# YOUR CODE HERE
```

```python
print("There are 88 airports in Illinois and you counted", num_apts_in_IL[0])
```

<!-- grade_id: p1_test, points: 5, grade: true, solution: false -->
```python
assert_equal(num_apts_in_IL, ['88'])
```

# Problem 2.

The main purpose of this problem is to extend what you learned in Problem 1 about data processing at the command line. Now you must list the airports in Illinois and store them in a file called `~/w4p1_dir/IL_APTS.txt`, with one airport per line. You already know how to find the lines corresponding to airports in Illinois (either by finding `"IL"` or `IL` in the state column or finding `"IL"` anywhere in the line), so now you need to extract the **second column** of those lines.

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
%%bash 
# YOUR CODE HERE
```

```python
res = !cat /home/data_scientist/w4p1_dir/IL_APTS.txt | sort
print("There are 88 airports in Illinois and you counted", len(res))
```

<!-- grade_id: p2_test, points: 10, grade: true, solution: false -->
```python
# insert assert tests here
res = !cat /home/data_scientist/w4p1_dir/IL_APTS.txt | sort
res_stripped = [x.strip('"') for x in res]
assert_equal(res_stripped, 
             ['Albertus',
             'Aurora Municipal',
             'Benton Municipal',
             'Cairo',
             'Campbell',
             'Capital',
             'Carmi Municipal',
             'Casey Municipal',
             'Centralia Municipal',
             'Central Illinois Regional',
             'Chanute',
             'Chicago Meigs',
             'Chicago Midway',
             "Chicago O'Hare International",
             'Clow',
             'Coles County Memorial',
             'Dacy',
             'Decatur',
             'Dekalb-Taylor Municipal',
             'Dixon Muni-Charles R Walgreen ',
             'Du Page',
             'Edgar County',
             'Effingham County Memorial',
             'Fairfield Municipal',
             'Flora',
             'Frankfort',
             'Galesburg Municipal',
             'Galt',
             'Gibson City Municipal',
             'Greater Beardstown',
             'Greater Kankakee',
             'Greater Peoria Regional',
             'Greater Rockford',
             'Greenville',
             'Harrisburg-Raleigh',
             'Havana Regional',
             'Illinois Valley Regional',
             'Ingersoll',
             'Jacksonville Municipal',
             'Joliet Park District',
             'Kewanee Municipal',
             'Lake In The Hills',
             'Lansing Municipal',
             'Lawrenceville-Vincennes Municipal',
             'Lewis University',
             'Litchfield Municipal',
             'Logan County',
             'Macomb Municipal',
             'Marshall County',
             'Mercer County',
             'Metropolis Municipal',
             'Monmouth Municipal',
             'Morris Municipal',
             'Mt Carmel Municipal',
             'Mt. Hawley Auxiliary',
             'Mt. Sterling Municipal',
             'Mt Vernon-Outland',
             'Olney-Noble',
             'Palwaukee',
             'Pekin Municipal',
             'Pinckneyville-Duquoin',
             'Pittsfield Penstone Municipal',
             'Pontiac Municipal',
             'Poplar Grove ',
             'Quad City',
             'Quincy Municipal-Baldwin ',
             'Robinson Municipal',
             'Rochelle Municipal',
             'Salem-Leckrone',
             'Sanger',
             'Schaumburg',
             'Schaumburg Heliport',
             'Scott AFB/MidAmerica ',
             'Shelby County',
             'Southern Illinois ',
             'Sparta Community-Hunter ',
             'St. Louis Downtown',
             'St Louis-Metro East',
             'St. Louis Regional',
             'Taylorville Municipal',
             'Tri-Township ',
             'Tuscola',
             'University of Illinois-Willard',
             'Vandalia Municipal',
             'Vermilion County',
             'Waukegan Regional',
             'Whiteside Co ',
             'Williamson County'])
```

# Problem 3.

Finally, store the Nothernmost and Southernmost airports in Illinois (as defined by latitude) in a text file called `~/w4p1_dir/N_S_IL.txt` with the Northernmost airport on the first line and the Southernmost airport on the second line. Keep in mind that latitude increases as you go North and is recorded in the 6th column of the text file. You may wish to find the Northernmost airport and write to an output file, then append the Southernmost airport to the end of the file with `>>`.

<!-- grade_id: p3_answer, points: 0, grade: false, solution: true -->
```python
%%bash 
# YOUR CODE HERE
```

```python
!echo your output is
!ec
!cat /home/data_scientist/w4p1_dir/N_S_IL.txt
```

<!-- grade_id: p3_tests, points: 10, grade: true, solution: false -->
```python
res3 = !cat /home/data_scientist/w4p1_dir/N_S_IL.txt
res3_stripped = [x.strip('"') for x in res3]
assert_equal(res3_stripped, ['Waukegan Regional', 'Cairo'])
```

## Cleaning Up

```python
!rm -rf /home/data_scientist/w4p1_dir/
```