# Week 4 Problem 3

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/info490-fa16/blob/master/Week2/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
from nose.tools import assert_equal
```

## Summary

This problem helps you get familiar with Python dictionaries. It is useful to refer to the dictionary functions mentioned in the [pydatastructures2.ipynb](https://github.com/UI-DataScience/info490-fa16/blob/master/Week4/notebooks/pydatastructures2.ipynb) notebook.

## 1. Function: dict_from_lists()

In the following code cell, construct a function named `dict_from_lists()` that takes two lists as input arguments. The function returns a dictionary with the first input list as its keys and the second as its values. You can assume that the two lists have equal lengths, and that the n-th element in the first list and the n-th element in the second list are one key-value pair.

<!-- grade_id: f1_ans, points: 0, grade: false, solution: true -->
```python
def dict_from_lists(keys, values):
    '''
    Constructs a dictionary using the first input list as keys and second as values.
    
    Parameters
    __________
    keys: A list.
    values: A list.
    
    Returns
    _______
    output: A dictionary.
    '''
    # YOUR CODE HERE
```

<!-- grade_id: f1_test, points: 5, grade: true, solution: false -->
```python
assert_equal(type(dict_from_lists([], [])), dict)
assert_equal(dict_from_lists([], []), {})
assert_equal(dict_from_lists([1, 2], ['1', '2']), {1: '1', 2: '2'})
assert_equal(dict_from_lists(tuple('MaRiO'), list(range(5))), 
             {'M': 0, 'a': 1, 'R': 2, 'i': 3, 'O': 4})
assert_equal(dict_from_lists(["info", "490", "fa16"], [4.0, 9.0, 0.0]), 
             {"info": 4.0, "490": 9.0, "fa16": 0.0})
```

## 2. Function: add_one()

In the following code cell, construct a function named `add_one()` that takes two dictionaries as input arguments, and adds the key-value pairs from the second dictionary to the first one. For example, suppose `a = {1: 2}` and `b = {3: 4}`, then `add_one(a, b)` returns `{1: 2, 3: 4}`. This function returns the updated first input dictionary as the output. 

Since Python dictionaries do not accept repeated key values, if a key in the second dictionary already exists in the first dictionary, your `add_one()` function should replace the corresponding value in the first dictionary with the value in the second dictionary. For example, suppose `a = {1: 2}` and `b = {1: 3}`, then `add_one(a, b)` returns `{1: 3}`. *Hint*: You do not need to write extra codes for duplicate keys. Python handles those automatically.

<!-- grade_id: f2_ans, points: 0, grade: false, solution: true -->
```python
def add_one(d, to_add):
    '''
    Adds the second dictionary to the first dictionary.
    
    Parameters
    __________
    d: A dictionary.
    to_add: A dictionary that will be added to d.
    
    Returns
    _______
    output: The updated dictionary d. 
    '''
    # YOUR CODE HERE
```

<!-- grade_id: f2_test, points: 5, grade: true, solution: false -->
```python
assert_equal(type(add_one({}, {})), dict)
assert_equal(add_one({}, {0: 0}), {0: 0})
assert_equal(add_one({1: 2, 3: 4}, {1: 3}), {1: 3, 3: 4})
assert_equal(add_one({1: '1', 2: '2'}, {'info': '490'}), {'info': '490', 1: '1', 2: '2'})
assert_equal(add_one({'M': 0, 'a': 1, 'R': 2}, {'i': 3, 'O': 4}), 
             {'M': 0, 'a': 1, 'R': 2, 'i': 3, 'O': 4})
assert_equal(add_one({"info": 4.0, "490": 9.0, "fa16": 0.0}, {'M': 0, 'a': 1, 'R': 2}), 
             {'M': 0, 'a': 1, 'R': 2, "info": 4.0, "490": 9.0, "fa16": 0.0})
```

## 3. Function: add_many():

In the following code cell, construct a function named `add_many()` that takes a dictionary as the first input argument, and arbitrary number of dictionaries as the second input argument. This function adds the key-value pairs from arbitrary number of dictionaries to the first input dictionary. For example, suppose `a = {1: 2}`, `b = {3: 4}`, and `c = {5: 6}`, then `add_many(a, b, c)` returns `{1: 2, 3: 4, 5: 6}`. For duplicate keys, update the corresponding values in the first dictionary; again, let Python handles those automatically. 

*Hint*: you can take advantage of the previously defined `add_one()` function to avoid duplicate codes.

<!-- grade_id: f3_ans, points: 0, grade: false, solution: true -->
```python
def add_many(d, *dicts_to_add):
    '''
    Adds many dictionaries to the first dictionary.
    
    Parameters
    __________
    d: A dictionary.
    *dicts_to_add: Arbitrary number of dictionaries (1 or more dictionaries) that will be added to d.
    
    Returns
    _______
    output: The updated dictionary d. 
    '''
    # YOUR CODE HERE
```

<!-- grade_id: f3_test, points: 5, grade: true, solution: false -->
```python
assert_equal(type(add_many({}, {}, {})), dict)
assert_equal(add_many({}, {0: 0}), {0: 0})
assert_equal(add_many({1: 2, 3: 4}, {1: 3}, {3: 5}), {1: 3, 3: 5})
assert_equal(add_many({1: '1', 2: '2'}, {3: '3'}, {4: '4', 5: '5'}), 
             {1: '1', 2: '2', 3: '3', 4: '4', 5: '5'})
assert_equal(add_many({'M': 0}, {'a': 1}, {'R': 2}, {'i': 3}, {'O': 4}), 
             {'M': 0, 'a': 1, 'R': 2, 'i': 3, 'O': 4})
assert_equal(add_many({"info": "4.0", "490": "9.0"}, {"fa16": "0.0"}, {'M': 0, 'a': 1, 'R': 2}), 
             {'M': 0, 'a': 1, 'R': 2, "info": "4.0", "490": "9.0", "fa16": "0.0"})
```

## 4. Function remove():

In the following code cell, construct a function named `remove()` that takes two dictionaries as the input arguments. When executed, this function removes all the key-value pairs in the second dictionary from the first dictionary *only if* those key-value pairs also exist in the first dictionary. The `remove` function returns the first dictionary as the output.

For example, suppose `a = {1: 2, 3: 4, 5: 6}`. 
- If `b = {1: 2}`, then `remove(a, b)` returns `{3: 4, 5: 6}`.
- If `b = {1: 2, 3: 4, 5: 6}`, then `remove(a, b)` returns `{}` (an empty dictionary).
- If `b = {0: 0}`, then `remove(a, b)` returns `{1: 2, 3: 4, 5: 6}`.
- If `b = {1: 0}` (key exists in dictionary `a` but the key-value pair does not), then `remove(a, b)` returns `{1: 2, 3: 4, 5: 6}`.

<!-- grade_id: f4_ans, points: 0, grade: false, solution: true -->
```python
def remove(d, to_remove):
    '''
    Removes the second dictionary from the first dictionary. 
    
    Parameters
    __________
    d: A dicionary.
    to_remove: A dictionary that will be removed from d. 
    
    Returns
    _______
    output: The updated dictionary d. 
    '''
    # YOUR CODE HERE
```

<!-- grade_id: f4_test, points: 5, grade: true, solution: false -->
```python
assert_equal(type(remove({1: 2, 3: 4}, {1: 2})), dict)
assert_equal(remove({1: '1', 2: '2', 3: '3', 4: '4', 5: '5'}, {1: '1'}), 
             {2: '2', 3: '3', 4: '4', 5: '5'})
assert_equal(remove({1: '1', 2: '2', 3: '3', 4: '4', 5: '5'}, {2: '2'}), 
             {1: '1', 3: '3', 4: '4', 5: '5'})
assert_equal(remove({1: '1', 2: '2', 3: '3', 4: '4', 5: '5'}, {3: '3', 4: '4'}), 
             {1: '1', 2: '2', 5: '5'})
assert_equal(remove({1: '1', 2: '2', 3: '3', 4: '4', 5: '5'}, {1: '1', 2: '2', 3: '3', 4: '4', 5: '5'}), 
             {})
assert_equal(remove({1: '1', 2: '2', 3: '3', 4: '4', 5: '5'}, {'a': 'b'}), 
             {1: '1', 2: '2', 3: '3', 4: '4', 5: '5'})
assert_equal(remove({1: '1', 2: '2', 3: '3', 4: '4', 5: '5'}, {1: 'a'}), 
             {1: '1', 2: '2', 3: '3', 4: '4', 5: '5'})
assert_equal(remove({1: '1', 2: '2', 3: '3', 4: '4', 5: '5'}, {1: 'a', 2: '2'}), 
             {1: '1', 3: '3', 4: '4', 5: '5'})
```

## 5. Eeveelution:

Before you work on this part, make sure all of the previous functions have passed the assertion tests. 

Now that we have `dict_from_lists()`, `add_one()`, `add_many()`, and `remove()` functions, let's try those functions on some data: 

```python
# lists
names = ["Eevee", "Vaporeon", "Jolteon", "Flareon", "Pikachu"]
types = ["Normal", "Water", "Electric", "Fire", "Electric"]

# dictionaries
generationII = {"Espeon": "Psychic", "Umbreon": "Dark"}
generationIV = {"Leafeon": "Grass", "Glaceon": "Ice"}
generationVI = {"Sylveon": "Fairy"}
```

In the code cell below, complete the following tasks:
- Construct a dictionary named `eeveelution` with the list `names` as keys and the list `types` as values.
- Add the dictionaries `generationII`, `generationIV`, `generationVI` to the dictionary `eeveelution`.
- Remove the key-value pair `("Pikachu", "Electric")` from the dictionary `eeveelution`.

Here is a sample output: 
```
{'Eevee': 'Normal', 'Espeon': 'Psychic', 'Flareon': 'Fire', 'Glaceon': 'Ice', 'Jolteon': 'Electric', 'Leafeon': 'Grass', 'Sylveon': 'Fairy','Umbreon': 'Dark','Vaporeon': 'Water'}
```

<!-- grade_id: f5_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: f5_test, points: 5, grade: true, solution: false -->
```python
assert_equal(type(eeveelution), dict)
assert_equal(len(eeveelution), 9)
assert_equal(eeveelution["Eevee"], "Normal")
assert_equal(eeveelution, {'Eevee': 'Normal', 'Espeon': 'Psychic', 'Flareon': 'Fire', 
                           'Glaceon': 'Ice', 'Jolteon': 'Electric', 'Leafeon': 'Grass',
                           'Sylveon': 'Fairy','Umbreon': 'Dark','Vaporeon': 'Water'})
```