# Week 6 Problem 1

If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/info490-fa16/blob/master/Week2/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
from nose.tools import assert_equal, assert_true, assert_false
```

```python
%%bash
# make sandbox directory
mkdir -p ~/w6p1_testdir
wget -O /home/data_scientist/w6p1_testdir/crossword.jpg https://fosterdisbelief.files.wordpress.com/2015/11/crossword1.jpg
cp /etc/bash.bashrc ~/w6p1_testdir
```

```python
from IPython.display import Image
Image(filename='/home/data_scientist/w6p1_testdir/crossword.jpg') 
```

# Problem 1.

The main purpose of this problem is to give you experience with the tools you learned last week in conjunction with the tools you learned this week. The crossword puzzle above has some entries missing. We'll use regular expressions to help solve a couple of the missing entries. Write a bash script to do the following: 

1. Go get a dictionary of english language words from [this link](http://www-01.sil.org/linguistics/wordlists/english/wordlist/wordsEn.txt)
2. Save the dictionary in ~/w6p1_testdir as `dictionary.txt`
3. 32 across (from the puzzle above) is a 4-letter word where the second letter is 'D'. Find all such words in the dictionary and save the output in ~/w6p1_testdir as `32_across.txt`.
4. The answer to 32 across is likely the word 'Edit'. 18 across is likely 'Ugly', which means 12 down is a 6-letter word of the form '??gre?'. Find all such words in the dictionary and save the output in ~/w6p1_testdir as `12_down.txt`

Ignore case as all words in the dictionary are lowercase.

**NOTE**: The `dictionary.txt` file includes a whitespace character trailing each line. You must consider this when constructing your regular expressions!

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
%%bash
#YOUR CODE HERE
```

<!-- grade_id: p1_test, points: 5, grade: true, solution: false -->
```python
# dictionary.txt is in the right place
list_dir = !ls /home/data_scientist/w6p1_testdir/ 
assert_true('dictionary.txt' in list_dir)

words_32_across = !cat /home/data_scientist/w6p1_testdir/32_across.txt
# 32_across.txt is the right length
assert_equal(len(words_32_across), 28)
# a few select items are in 32_across.txt
assert_true('adam' in words_32_across)
assert_true('edit' in words_32_across)
assert_true('idea' in words_32_across)

words_12_down = !cat /home/data_scientist/w6p1_testdir/12_down.txt
# 32_across.txt is the right length
assert_equal(len(words_12_down), 3)
# a few select items are in 32_across.txt
assert_true('degree' in words_12_down)
assert_true('regret' in words_12_down)
assert_true('regrew' in words_12_down)
```

# Problem 2.

The main purpose of this problem is to give you experience constructing slightly more complicated regular expressions for a real-world purpose. There is a file in `/etc` called `bash.bashrc` that should be copied into your ~/w6p1_testdir. Suppose you want another copy of it, but you don't want any of the comments that come with it. Create two files in ~/w6p1_testdir:

1. All the commented lines in `/etc/bash.bashrc` as bashrc_comments.txt
2. All the non-commented lines **that are not whitespace** as bashrc_commands.txt

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
%%bash
#YOUR CODE HERE
```

<!-- grade_id: p2_test, points: 10, grade: true, solution: false -->
```python
# is everything where it's supposed to be?
files_in_testdir = !ls /home/data_scientist/w6p1_testdir/
assert_true('bashrc_comments.txt' in files_in_testdir)
assert_true('bashrc_commands.txt' in files_in_testdir)

# are the comments properly extracted?
comments = !cat /home/data_scientist/w6p1_testdir/bashrc_comments.txt
assert_equal(len(comments), 27)

# are the commands properly extracted
commands = !cat /home/data_scientist/w6p1_testdir/bashrc_commands.txt
assert_false("# check because c-n-f could've been removed in the meantime" in [x.strip() for x in commands])
assert_equal(len(commands), 20)
```

### Bonus
What if we wanted to run the commands in the newly created script? We can do so with the following:

```python
%%bash
cd ~/w6p1_testdir/
chmod u+x bashrc_commands.txt
source bashrc_commands.txt
```

# Problem 3.

The main purpose of this problem is to give you more exposure to constructing regular expressions. We return to the `bash.bashrc` file but this time we want to find all the conditional (`if` or `elif`) statements. Specifically, make the following inside of ~/w6p1_testdir:

1. A file `conditionals.txt` that contains all the `if` and `elif` statements in the `/etc/bash.bashrc` file.
2. A file `commented_conditionals.txt` that contains all the `if` and `elif` statements in `/etc/bash.bashrc` that have been commented out

You may wish to use the `-E` flag with `grep`.

<!-- grade_id: p3_answer, points: 0, grade: false, solution: true -->
```python
%%bash
#YOUR CODE HERE
```

<!-- grade_id: p3_test, points: 10, grade: true, solution: false -->
```python
# did you find all the conditionals?
conditionals = !cat /home/data_scientist/w6p1_testdir/conditionals.txt
assert_equal(len(conditionals), 4)
assert_true("elif [ -x /usr/share/command-not-found/command-not-found ]; then" in [x.strip() for x in conditionals])

# did you find the correct commented conditionals?
commented_conditionals = !cat /home/data_scientist/w6p1_testdir/commented_conditionals.txt
assert_equal(len(commented_conditionals), 3)
assert_false("# if the command-not-found package is installed, use it" in [x.strip() for x in commented_conditionals]) 
```

## Cleaning Up

```python
rm -rf /home/data_scientist/w6p1_testdir/
```

## Just in Case

If you accidentally blow up your /etc/bash.bashrc, here it is:

```
# System-wide .bashrc file for interactive bash(1) shells.

# To enable the settings / commands in this file for login shells as well,
# this file has to be sourced in /etc/profile.

# If not running interactively, don't do anything
[ -z "$PS1" ] && return

# check the window size after each command and, if necessary,
# update the values of LINES and COLUMNS.
shopt -s checkwinsize

# set variable identifying the chroot you work in (used in the prompt below)
if [ -z "${debian_chroot:-}" ] && [ -r /etc/debian_chroot ]; then
    debian_chroot=$(cat /etc/debian_chroot)
fi

# set a fancy prompt (non-color, overwrite the one in /etc/profile)
PS1='${debian_chroot:+($debian_chroot)}\u@\h:\w\$ '

# Commented out, don't overwrite xterm -T "title" -n "icontitle" by default.
# If this is an xterm set the title to user@host:dir
#case "$TERM" in
#xterm*|rxvt*)
#    PROMPT_COMMAND='echo -ne "\033]0;${USER}@${HOSTNAME}: ${PWD}\007"'
#    ;;
#*)
#    ;;
#esac

# enable bash completion in interactive shells
#if ! shopt -oq posix; then
#  if [ -f /usr/share/bash-completion/bash_completion ]; then
#    . /usr/share/bash-completion/bash_completion
#  elif [ -f /etc/bash_completion ]; then
#    . /etc/bash_completion
#  fi
#fi

# if the command-not-found package is installed, use it
if [ -x /usr/lib/command-not-found -o -x /usr/share/command-not-found/command-not-found ]; then
	function command_not_found_handle {
	        # check because c-n-f could've been removed in the meantime
                if [ -x /usr/lib/command-not-found ]; then
		   /usr/lib/command-not-found -- "$1"
                   return $?
                elif [ -x /usr/share/command-not-found/command-not-found ]; then
		   /usr/share/command-not-found/command-not-found -- "$1"
                   return $?
		else
		   printf "%s: command not found\n" "$1" >&2
		   return 127
		fi
	}
fi
```